Age and Learning Shapes Sound Representations in Auditory Cortex During Adolescence

Author response:

The following is the authors’ response to the previous reviews.

Reviewer #1 (Public review):

Summary:

Praegel et al. explore the differences in learning an auditory discrimination task between adolescent and adult mice. Using freely-moving (Educage) and head-fixed paradigms, they compare behavioral performance and neuronal responses over the course of learning. The mice were initially trained for seven days on an easy pure frequency tone Go/No-go task (frequency difference of one octave), followed by seven days of a harder version (frequency difference of 0.25 octave). While adolescents and adults showed similar performance on the easy task, adults performed significantly better on the harder task. Quantifying the lick bias of both groups, the authors then argue that the difference in performance is not due to a difference in perception, but rather to a difference in cognitive control. The authors then used neuropixel recordings across 4 auditory cortical regions to quantify the neuronal activity related to the behavior. At the single cell level, the data shows earlier stimulus-related discrimination for adults compared to adolescents in both the easy and hard tasks. At the neuronal population level, adults displayed a higher decoding accuracy and lower onset latency in the hard task as compared to adolescents. Such differences were not only due to learning, but also to age as concluded from recordings in novice mice. After learning, neuronal tuning properties had changed in adults but not in adolescent. Overall, the differences between adolescent and adult neuronal data correlates with the behavior results in showing that learning a difficult task is more challenging for younger mice.

Strengths:

The behavioral task is well designed, with the comparison of easy and difficult tasks allowing for a refined conclusion regarding learning across age. The experiments with optogenetics and novice mice are completing the research question in a convincing way.

The analysis, including the systematic comparison of task performance across the two age groups, is most interesting, and reveals differences in learning (or learning strategies?) that are compelling.

Neuronal recording during both behavioral training and passive sound exposure is particularly powerful, and allows interesting conclusions.

Weaknesses:

The presentation of the paper must be strengthened. Inconsistencies, missing information or confusing descriptions should be fixed.

We have carefully re-read the manuscript and reviewed it for inconsistencies. We made several corrections in the figures. For example, we removed redundant lines from violin plots and statistics, applied consistent labels, matched y- and x-limits of graphics, and adjusted labels. We also clarified descriptions of some experiment by adding explanations to the text.

The recording electrodes cover regions in the primary and secondary cortices. It is well known that these two regions process sounds quite differently (for example, one has tonotopy, the other not), and separating recordings from both regions is important to conclude anything about sound representations. The authors show that the conclusions are the same across regions for Figure 4, but is it also the case for the subsequent analysis? Comparing to the original manuscript, the authors have now done the analysis for AuDp and AUDv separately, and say that the differences are similar in both regions. The data however shows that this is not the case (Fig S7). And even if it were the case, how would it compatible with the published literature?

To address this and previous concerns about regional differences, the manuscript now includes 4 figures (4-1, 4-3, 6-2, 7-1) and 5 supplemental tables (3,4, 5, 6, 8) that explicitly compare results across brain regions.

Following the reviewer’s request for subsequent analysis, we now added a new supplemental figure (Fig. S6-2) and two new supplementary tables (Tables S5, S6). We show that similar to expert mice (supplementary Table 3, and supplementary Table 4), the firing properties of adolescent and adult novice mice differ across auditory subregions (supplementary Table 5). We also show that the different auditory subregions have different firing properties (supplementary Table 6). With respect to task engagement, we show that (similar to Fig. S4-2) the neuronal discriminability in different auditory subregions is similar in both novice and expert mice (Fig. S6-2).

Following the comment on Fig. S7-1, we made three changes to the revised manuscript. First, we now highlight that the differences firing properties between adolescent and adult neurons in AUDp and AUDv were distinct, but not significantly different within age-group comparisons. Second, we clearly state that the learning related changes in the measured parameters are different between AUDp and AUDv. Note, however, the greater changes in adult neurons after learning remains consistent between AUDp and AUDv. Third, we softened our original claim but still highlighted the stronger learning-induced plasticity in adults.

Regarding the concern that different regions should show different patterns due to their known differences (e.g. tonotopy). Of course we agree that different areas differ functionally (as shown in our own previous work and here as well). However, it is still plausible, and biologically reasonable, that developmental changes may proceed in a similar direction across different areas, even if their baseline coding properties differ.

Reviewer #2 (Public review):

Summary:

The authors aimed to find out how and how well adult and adolescent mice discriminate tones of different frequencies and whether there are differences in processing at the level of the auditory cortex that might explain differences in behavior between the two groups. Adolescent mice were found to be worse at sound frequency discrimination than adult mice. The performance difference between the groups was most pronounced when the sounds are close in frequency and thus difficult to distinguish and could, at least in part, be attributed to the younger mice' inability to withhold licking in no-go trials. By recording the activity of individual neurons in the auditory cortex when mice performed the task or were passively listening as well as in untrained mice the authors identified differences in the way that the adult and adolescent brains encode sounds and the animals' choice that could potentially contribute to the differences in behavior.

Strengths:

The study combines behavioural testing in freely-moving and head-fixed mice, optogenetic manipulation and high density electrophysiological recordings in behaving mice to address important open questions about age differences in sound-guided behavior and sound representation in the auditory cortex.

Weaknesses:

For some of the analyses that the authors conducted it is unclear what the rationale behind them is and, consequently, what conclusion we can draw from them.

We have carefully re-read the manuscript and reviewed it for analyses that lacked a clear rationale or conclusion. To address this, we have made several changes to clarify the reasoning and strengthen the interpretation of the results.

Reviewer #1 (Recommendations for the authors):

It would have helped if the authors had highlighted the changes they made to the manuscript compared to the original version - especially since many replies to the reviewers' comments were as vague as "...we fixed some of the wording so it adheres to the data shown", or "we refined our interpretation", without further details.

The revised version has improved substantially, and the main claims have been discussed in a more objective way. Important new analyses have been added to allow for a refined interpretation of the results. However, the presentation of the data could still be strengthened significantly (in response to comment A from last review).

We apologize for the lack of detail in some of our previous responses. Our intention was to keep the replies concise, assuming that the side-by-side version with tracked changes would make the edits sufficiently clear. However, we understand the need for greater transparency. Thus, below we provide the following five lists describing the major changes: (1) List of specific reviewer recommendations, (2) list of corrections in figures, (3) list of clarity issues, (4) list of fixed mistakes, (5) list of new figures. We hope this breakdown makes the revisions clearer and more accessible.

List of specific reviewer recommendations:

l.108 mentions a significant change in the vertical line of Fig 1F - Could this significance be indicated and quantified in the figure?

We quantified and indicated the significance of the vertical line in Fig. 1f and Fig. 1i.

Fig.1G - the thick and thin lines should be defined, as well as the grey and white dots (same values for adolescents, not for adults).

(a) We removed the thin inner lines from the violin plot. We define the bar (thick line) of the violin plot in an additional sentence in the methods section under data analysis (LL820-823). b) We adjusted the marker outlines in the adult data (Fig. 1G).

the figure axis legends should be consistent (trails in Fig D vs # trails in Fig 1F)

We adjusted the axis legend to # trials in Fig. 1D.

l.110: is d' always calculated based on the 100 last trials of a session, or is it just for Figure 1F? -etc...

d’ is always calculated based on the last 100 trials. To clarify this, we added a description in the methods section (L830).

List of corrections in the figures:

(1) We removed the internal lines from violin plots in throughout Fig. 1-7.

(2) We removed the underline of the statistics throughout Fig. 1-7.

(3) We consistently applied ‘adolescent’ and ‘adult’ figure labels and titles with lowercase letters throughout Fig. 1-7.

(4) We applied consistent labelling of ‘time (ms)’ throughout Fig. 1-7.

(5) We matched the size of dashed lines throughout Fig. 1-6.

(6) We adjusted the x-label of Fig. 1d, Fig. S-1-1 a, Fig. 3c, Fig. 3h-i, Fig, 4d to ‘# trials’.

(7) We removed the x-label of ‘Experimental Group’ from Fig. 1 to enhance consistency with other figures.

(8) We removed misaligned dots from the violin plots in Fig. 1g, Fig. 2f, Fig. 3f,g.

(9) We corrected the plot in Fig. S1-1b.

(10) We adjusted the y-limits of Fig. S1-1c to be consistent with Fig. S1-1d,e.

(11) We adjusted the x-labels and y-labels of Fig, 2, Fig. S3-1, Fig, S3-2 and Fig. 3b to ‘freq. (kHz)’.

(12) We added the age of adolescent and adult mice to the schematic timeline in Fig. 2a.

(13) We added a label of the reinforcement delay to the schematic trial structure in Fig. 3b.

(14) We added within-group statistics to Fig. 3e and the figure legend.

(15) We adjusted the x-label of Fig. 3d to ‘# sessions’.

(16) We adjusted the x-label of Fig. 3d and Fig. S3-1b to ‘# licks’.

(17) We changed the y-label in Fig. S3-1a, and Fig. S3-2d, e to ‘lick ratio’ to avoid confusion with the lick rate (Hz) that was calculated in Fig. 4 and Fig. 6.

(18) We replaced the titles ‘CAMKII’ with ‘dTomato’ in Fig. S3-2 to correctly highlight that both the experimental and control injection were CAMKII injections.

(19) We adjusted the x-labels and y-labels of Fig, 2, Fig. S3-1, Fig, S3-2 and Fig. 3b to ‘freq. (kHz)’.

(20) We adjusted the y-label of Fig. S4-1c to ‘# neurons’.

(21) We matched the x-ticks in Fig. 4e,f.

(22) We matched the x-ticks in Fig. 6d-g.

(23) We changed the x-label in Fig. 4g, S4-2 and S6-2 to ‘duration (ms)’ to match the figure label with the manuscript.

(24) We consistently label ‘Hit’, ‘Miss’, ‘FA’ and ‘CR’ with capital letters in Fig. 4d-e.

(25) We replaced the double figure label ‘C.’ in Fig. S4-2 with ‘D.’.

(26) We adjusted the dot-size in Fig. 5 to be equal for all graphs.

(27) We added ticks to the experimental timeline in Fig. 6a.

(28) We corrected the y-label in Fig.7c. Now it correctly reflects 5 attenuations from 72-32 dB SPL.

(29) We matched the y-label of Fig. 7e-h and Fig. S7-1.

List of clarity issues:

(1) We replaced the term ‘lower response bias’ with ‘higher lick bias’ (L24) to accurately describe the more negative (lower) criterion-bias, which highlights a higher tendency to lick.

(2) We replaced the term ‘response bias’ with ‘lick bias’ to consistently describe the calculated criterion-bias (L24, L149, L164, L455, L456, L468).

(3) We clarify that the age-related differences were ‘more pronounced’ instead of simply ‘higher’ to accurately reflect not simply the increase in adolescent lick-bias, but also the decrease in adult lick-bias (L31).

(4) We clarified that adolescent sound representations are not merely ’distinct’, but ‘not fully mature’ in L83.

(5) We clarified in L180 that the impulsive responses we observed in adolescent mice could be related to being ‘less impacted by punishments’.

(6) We clarified the differences in firing properties of auditory sub-regions analyzed in Supplementary Table 3 (L287-295).

(7) We explained and clarified the reference to Fig. 3j (LL252-253).

(8) We added statistics to Fig.S4-2 to support our claim that there are no differences in the onset-latency, duration of discriminability and maximal discriminability between different sub-regions within age-groups (LL 314-315).

(9) We expanded our explanation of the results in Table 3 (LL370-379).

(10) We separated the reference to Fig. 6b and Fig. 6c to clarify their meaning (LL358-361).

(11) We clarified the differences in basic firing properties during the FRA protocol in Fig. 7 (LL409-418).

(12) We expanded our explanation of the differences of the learning related firing properties in AUDp and AUDv of Fig. S7-1 (LL426-433).

(13) We changed the term ‘plasticity profiles’ to ‘learning related plasticity’ to further clarify our limitation that L5/6 and L2/3 may exhibit distinct learning related changes (L496).

(14) We changed the term ‘sluggish’ (L481) to ‘delayed’ to more precisely explain differences between adolescent and adult tuning properties.

(15) We clarified that the running d’ was calculated in bins of 25 trials, instead of ‘the last 25 trials’ (LL845-846).

List of fixed mistakes:

(1) We corrected and matched the age to more accurately reflect the age mice were recorded (P37-42 and P77-82).

(2) We corrected the attenuation range from 72-42 to 72-32 dB SPL to correctly reflect the 5 attenuations used in the protocol.

(3) We corrected the number of channels shown in the voltage trace from 10 to 11 (Fig. S4-1a)

(4) We corrected the number of neurons recorded in novice adolescent mice in the legend of Fig. 6 from 140 to 130 (Fig. 6b).

(5) We removed redundant, or double brackets, commas, dots, and semi-colons in the figure legends.

(6) We corrected the LME statistics Table 2.

List of new figures and tables:

(1) We added a new supplementary figure to accompany Figure 6. Specifically, Fig. S6-2, shows the interaction of the three measured discriminability properties (onset delay, duration of discriminability, and maximal discriminability) in novice compared to expert mice in the easy and hard task (Go compared to No Go). The figure compares the different auditory sub-regions (similar to Fig. S4-2). We show that the discriminability properties within different groups is not significantly different among the four different sub-regions.

(2) Supplementary Table 5: We compared the firing properties in different auditory subregions in novice mice, and found (similar to expert mice) that the firing properties differ between adult and adolescent mice across the four different sub-regions.

(3) Supplementary Table 6: We compared the firing properties between different subregions, separately for adolescent and adult novice mice. Similar to expert mice, we found that different auditory subregions differ in their auditory firing properties.

Reviewer #2 (Recommendations for the authors):

The authors largely addressed my suggestions.

Comparing hit vs correct rejection trials in the population decoding analysis (L313-314): The authors acknowledge that comparing these two trial types conflates choice and stimulus decoding but I am not convinced that the changes to the manuscript text make this clear enough to the reader.

Thank you for pointing this out. We have made additional revisions to clarify this, and other issues more explicitly, as follows:

(1) We have expanded the explanation of how our population decoding analysis conflates stimulus and choice, and we acknowledge the limitations of this approach in the Abstract (L28), the Results section (L324-326, LL367-370) and the Discussion (LL516-519).

(2) We replaced the analysis of impulsivity on the head-fixed task. Instead of analyzing all it is, we focus only on ITIs following FA trials (Fig. S3-1c,d). This is more consistent with the analysis in the Educage (Fig. S2-1), where we show that adolescents exhibit increased impulsivity after FA trials. We found a similar result for ITIs following FA trials in the head-fixed task.

(3) To provide complementary insight, we now further justify our use of the Fisher separation metric alongside decoding accuracy in Figure 5, with a clearer rationale provided in LL343-345

(4) We also clarified our reasoning for focusing on 62 dB SPL in the FRA-based analysis in LL400-403.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Adolescence is a developmental period characterized by heightened plasticity, yet how ongoing development affects sensory processing and cognitive function is unclear. We investigated how adolescent (postnatal day 20-40) and adult (postnatal day 60-80) mice differ in performance on a pure tone Go/No-Go auditory discrimination task of varying difficulty. Using dense electrophysiological recordings, we measured spiking activity at single neuron resolution in the auditory cortex while mice were engaged in the task. As compared to adults, adolescent mice showed lower auditory discrimination performance in the harder task. This performance difference was due to higher response variability and weaker cognitive control expressed as lower response bias. Adolescent and adult neuronal responses differed only slightly in representations of pure tones when measured outside the context of learning and the task. However, cortical representations after learning within the context of the task were markedly different. We found differences in stimulus- and choice-related activity at the single neuron level representations, as well as lower population-level decoding in adolescents. Overall, cortical decoding in adolescents was lower and slower, especially for difficult sound discrimination, reflecting immature cortical representations of sounds and choices. Notably, we found age-related differences, which were higher after learning, reflecting the combined impact of age and learning. Our findings highlight distinct neurophysiological and behavioral profiles in adolescence, underscoring the ongoing development of cognitive control mechanisms and cortical plasticity during this sensitive developmental period.

Revised version: Adolescence is a developmental period characterized by heightened plasticity. Yet, how ongoing development affects sensory processing and cognitive function is unclear. We investigated how adolescent (postnatal day 20-42) and adult (postnatal day 60-82) mice differ in performance on a pure tone Go/No-Go auditory discrimination task of varying difficulty. Using dense electrophysiological recordings, we measured spiking activity at single neuron resolution in the auditory cortex while mice were engaged in the task. As compared to adults, adolescent mice showed lower auditory discrimination performance in a difficult task. This difference in performance was due to higher response variability and weaker cognitive control expressed as higher lick bias. Adolescent and adult neuronal responses differed only slightly in representations of pure tones when measured outside the context of learning and the task. However, cortical representations after learning within the context of the task were markedly different. We found differences in stimulus- and choice-related activity at the single neuron level representations, as well as lower population-level decoding of the difficult task in adolescents. Overall, cortical decoding in adolescents was lower and slower, especially for difficult sound discrimination, reflecting immature cortical representations of sounds and choices. Notably, we found age-related differences, which were more pronounced after learning, reflecting the combined impact of age and learning. Our findings highlight distinct neurophysiological and behavioral profiles in adolescence, underscoring the ongoing development of cognitive control mechanisms and cortical plasticity during this sensitive developmental period.

Introduction

Earlier version: Adolescence is a developmental stage characterized by the continued refinement of perception and cognition. During this period, the brain is highly sensitive to environmental influences, making it particularly vulnerable to both negative experiences, such as substance abuse, and positive influences, such as supportive relationships[1, 2]. While the adolescent brain exhibits greater plasticity compared to adults[3], the extent to which specific sensory and cognitive traits are fully developed versus those that remain malleable is still debated. For instance, studies on sensory perception have demonstrated that adolescent learning is slower and more variable[4, 5], while others have found that adolescents learn faster in novel contexts, such as reversal learning[6]. Certain traits, such as cognitive flexibility, are more pronounced in adolescents compared to adults, whereas other behaviors such as impulsive control, are still immature[7]. Furthermore, results across studies that target sensory or cognitive traits are often inconsistent[8]. Therefore, understanding the degree to which sensory and cognitive functions have matured and the underlying brain mechanisms that support them remain an ongoing challenge.

Revised version: Adolescence is a developmental stage characterized by the continued refinement of perception and cognition. During this period, the brain is highly sensitive to environmental influences, making it particularly vulnerable to both negative experiences, such as substance abuse, and positive influences, such as supportive relationships[1, 2]. While the adolescent brain exhibits greater plasticity compared to adults[3], the extent to which specific sensory and cognitive traits are fully developed versus those that remain malleable is still debated. For instance, some studies on sensory perception have demonstrated that adolescent learning is slower and more variable[4, 5], while others have found that adolescents learn faster in novel contexts, such as reversal learning[6]. Certain traits, such as cognitive flexibility, are more pronounced in adolescents compared to adults, whereas other behaviors such as impulsive control, are still immature[7]. Furthermore, results across studies that target sensory or cognitive traits are often inconsistent[8]. Therefore, understanding the degree to which sensory and cognitive functions have matured, and the underlying brain mechanisms that support them, remain an ongoing challenge.

A key concept in brain development is that of critical periods (CPs), which are specific time windows during which the brain’s neural circuits are highly plastic and particularly sensitive to external experiences[9, 10]. The early postnatal development of sensory cortices provides classical examples of this phenomenon. In mice, neurons in the visual cortex undergo several CPs such as those for orientation selectivity, direction selectivity and ocular dominance. During the CP, which occurs in different developmental time windows for different functional attributes, visual experience shapes the specific neuronal attribute to a long-term state (e.g., ocular dominance reaches its adult form)[9, 11]. The CPs for simple neuronal features in primary sensory cortices are typically complete by the onset of puberty. In the auditory cortex (ACx) of mice, the CP for pure tone processing begins as early as postnatal day 12 (P12), shortly after the ear canals open, and closes by ∼P15[12, 13]. Thus, the closure of the auditory CP to pure tones occurs before adolescence begins[14, 15]. Yet, while some studies have found that simple neuronal properties in the ACx (like response properties to pure tones) are nearly mature by adolescence, others report continued maturation well beyond this developmental stage[16, 17]. For example, previous studies found that auditory learning and behavioral performance in auditory discrimination tasks involving amplitude modulation detection, as well as temporal interval discrimination tasks remain highly constrained during adolescence[4, 5]. In those tasks, ACx neurons in adolescents exhibit high neuronal variability and lower tone sensitivity as compared to adults[4].

Earlier version: The perspective shift from viewing development as involving a set of discrete CPs to a broader view of continuous development with multiple, overlapping, CPs that affect each other, calls for studies to test the extent of functional modulation in neuronal features after their CP ‘closes’. Thus, although the CP for pure tone representation in the ACx ends by ∼P15, we hypothesized that learning of pure tone discrimination would remain malleable in adolescence, because other perceptual features and particularly other cognitive features are still developing. Further, it remains unclear how these features are expressed when measured in the context of active behavior. We measured cortical activity in the ACx of adolescent mice engaged in an auditory discrimination task involving pure tones. Our findings show that learning, behavior, and sound representations to pure tones in the ACx are distinct in adolescence.

Revised version: The perspective shift from viewing development as involving a set of discrete CPs to a broader view of continuous development with multiple, overlapping, CPs that affect each other, calls for studies to test the extent of functional modulation in neuronal features after their CP ‘closes’. Thus, although the CP for pure tone representation in the ACx ends by ∼P15, we hypothesized that learning of pure tone discrimination would remain malleable in adolescence, because other perceptual features and particularly other cognitive features are still developing. Further, it remains unclear how these features are expressed when measured in the context of active behavior. We measured cortical activity in the ACx of adolescent mice engaged in an auditory discrimination task involving pure tones. Our findings show that learning, behavior, and sound representations to pure tones in the ACx are not fully mature during adolescence.

Results

Adult mice outperform adolescents on a difficult, but not an easy, auditory discrimination task of pure tones

Earlier version: To study learning and perceptual performance abilities we trained and tested adolescent and adult mice on an auditory pure-tone discrimination Go/No-Go paradigm. For this purpose, we utilized an automated behavioral platform called the ‘Educage’ (Fig.1a)[18]. We trained adolescent (n = 15) and adult (n = 15) mice from post-natal day 20 (P20)-P37 and P60-P77, respectively (Fig. 1b). Following a tone association period (3-days; exposure to Go tones when mice enter the drinking port), mice learned to discriminate between two pure tones separated by 1 octave (7.07kHz vs 14.14kHz; Fig. 1b-c, light blue; herein referred to as the ‘easy task’). After one week of training on the easy task (P23 to P30 in pre-adolescent mice and P63 to P70 in adult mice) mice were introduced to a second pure-tone pair, now separated by 0.25 octave (9.2kHz vs 10.9kHz) (Fig. 1b-c, dark blue; herein referred to as the ‘hard task’). We completed the experiment after 1 week of learning on the easy and hard task simultaneously (i.e., at P37 for adolescents and P77 for adults). During the whole period of training, we also presented mice with low-probability, unrewarded ‘probe’ trials of pure-tones spanning the range of learned stimuli (Fig. 1c, grey). Figure 1d shows example learning curves (d’) from one adolescent and one adult mouse over the 14-day course of tone discrimination. Adolescents and adults learned the procedure of the task similarly well, as measured by the number of trials it took mice to first reach a discriminability threshold of d’ value ≥ 1 (Fig. 1e).

Revised version: To study learning and perceptual performance abilities we trained and tested adolescent and adult mice on an auditory pure-tone discrimination Go/No-Go paradigm. For this purpose, we utilized an automated behavioral platform called the ‘Educage’ (Fig. 1a)[18]. We trained adolescent (n = 15) and adult (n = 15) mice from post-natal day 20 (P20)-P37 and P60-P77, respectively (Fig. 1b). Following a tone association period (3-days; exposure to Go tones when mice enter the drinking port), mice learned to discriminate between two pure tones separated by 1 octave (7.07kHz vs 14.14kHz; Fig. 1b-c, light blue; herein referred to as the ‘easy task’). After one week of training on the easy task (P23 to P30 in pre-adolescent mice and P63 to P70 in adult mice) mice were introduced to a second pure-tone pair, now separated by 0.25 octave (9.2kHz vs 10.9kHz; Fig. 1b-c, dark blue; herein referred to as the ‘hard task’). We completed the experiment after 1 week of learning on the easy and hard task simultaneously (i.e., at P37 for adolescents and P77 for adults). During the whole period of training, we also presented mice with low-probability, unrewarded ‘probe’ trials of pure-tones spanning the range of learned stimuli (Fig. 1c, grey). Figure 1d shows example learning curves (of d’) from one adolescent and one adult mouse over the 14-day course of tone discrimination. Adolescents and adults learned the procedure of the task similarly well, as measured by the number of trials it took mice to reach a discriminability threshold of d’ value ≥ 1 (Fig. 1e).

Adolescent mice exhibit lower performance during self-initiated auditory learning in the ‘Educage’.

Earlier version: A. Schematic model of the ‘Educage’ (left), the trial structure and trial types (FA, false alarm; CR, correct reject). Created with BioRender.com. B. Experimental timeline. Total training time was 21 days (±2). C. The sounds used for training. Light blue-easy task; dark blue-hard task; Grey-probe-trials. D. Learning curve examples. Adolescent mouse (left, P20-to-P37); adult mouse (right, P60-to-P77). Vertical dashed lines indicate the easy-hard transition. Horizontal line is d’=1. E. Number of trials to reach threshold (d’>=1; adolescents, n =15; adults, n=15; z=-0.1659; p=0.8682, two-sample Wilcoxon rank sum test). F. Discriminability (d’) of the easy task in adolescent mice at P30 (grey; n=15, d’=2.0001±0.1791; n trials=8317±712) and adult mice at P70 (black; n=15, d’=2.1986 ± 0.2441; n trials=13229 ± 514, top: marked by the arrowhead). Dashed lines: mean trials per group (t-stat=-5.6314, df=28 ,p=4.9566e-06, two-sample independent t-test, vertical line), and mean d’ per age group (z=-0.2074; p=0.8357, two-sample Wilcoxon rank sum test, horizontal line) G. Change in discriminability (Δd’) of the easy task before and after the introduction of the hard task (Top: arrowheads; left: adolescent, signed rank=14, p=0.0067; right: adult, signed rank=14, p=0.1514, one-sample Wilcoxon signed rank test; Δd’ between adult and adolescent mice: z=-2.0739; p=0.0381, two-sample Wilcoxon rank sum test) H. Same as ‘G’ for the first 100 and last 100 trials of the experiment in the easy task (adolescent signed rank=120, p=6.1035e-05; adult signed rank=120, p=6.1035e-05, one-sample Wilcoxon signed rank test; Δd’ between adult and adolescent mice: z=-1.0370; p=0.2998, two-sample Wilcoxon rank sum test). I. Same as ‘F’ for the hard task (adolescent-grey; n=15, d’=1.1895±0.1783; n trials=4163 ± 297; adult-black; n=15, d’=1.8342 ± 0.1743; n trials=4102±475) (mean trials per group: t-stat=0.1306, df=28, p=0.8970, two-sample independent t-test, vertical line; mean d’ per group: z=-2.2398; p=0.0251, two-sample Wilcoxon rank sum test, horizontal line). J. Same as ‘G’ for the hard task. (adolescent signed rank=73, p=0.4887; adult signed rank=114, p=8.5449e-04, one-sample Wilcoxon signed rank test; Δd’ between adult and adolescent mice: z=-1.9495; p=0.0512, two-sample Wilcoxon rank sum test).

Revised version: A. Schematic model of the ‘Educage’ (left), the trial structure and trial types (FA, false alarm; CR, correct reject). Created with BioRender.com. B. Experimental timeline. Total training time was 21 days (±2). C. The sounds used for training. Light blue: easy task; Dark blue: hard task; Grey: catch-trials. D. Learning curve examples. Adolescent mouse (left, P20-to-P37); adult mouse (right, P60-to-P77). Vertical dashed lines indicate the easy-hard transition. Horizontal line is d’=1. E. Number of trials to reach threshold (d’>=1; adolescents, n = 15; adults, n = 15; z = −0.1659; p = 0.8682, two-sample Wilcoxon rank sum test). F. Discriminability (d’) of the easy task in adolescent mice at P30 (grey; n = 15, d’ = 2.0001 ± 0.1791; # trials = 8317 ± 712) and adult mice at P70 (black; n=15, d’=2.1986 ± 0.2441; # trials = 13229 ± 514, top: marked by the arrowhead). Dashed lines: mean trials per group (t-stat = −5.6314, df = 28, p = 4.9566e-06, two-sample independent t-test, solid vertical line), and mean d’ per age group (z = 0.2074; p = 0.8357, two-sample Wilcoxon rank sum test, solid horizontal line) G. Change in discriminability (Δd’) of the easy task before and after the introduction of the hard task (Top: arrowheads; left: adolescent, signed rank = 14, p = 0.0067; right: adult, signed rank = 14, p = 0.1514, one-sample Wilcoxon signed rank test; Δd’ between adult and adolescent mice: z = - 2.0739; p = 0.0381, two-sample Wilcoxon rank sum test) H. Same as ‘G’ for the first 100 and last 100 trials of the experiment in the easy task (adolescent signed rank = 120, p = 6.1035e-05; adult signed rank = 120, p = 6.1035e-05, one-sample Wilcoxon signed rank test; Δd’ between adult and adolescent mice: z = −1.0370; p = 0.2998, two-sample Wilcoxon rank sum test). I. Same as ‘F’ for the hard task (adolescent-grey; n =15, d’ = 1.1895 ± 0.1783; # trials=4163 ± 297; adult-black; n=15, d’=1.8342 ± 0.1743; # trials=4102±475; mean trials per group: t-stat = 0.1306, df = 28, p = 0.8970, two-sample independent t-test, solid vertical line; mean d’ per group: z = −2.2398; p = 0.0251, two-sample Wilcoxon rank sum test, solid horizontal line). J. Same as ‘G’ for the hard task. (adolescent signed rank = 73, p = 0.4887; adult signed rank = 114, p = 8.5449e-04, one-sample Wilcoxon signed rank test; Δd’ between adult and adolescent mice: z = −1.9495; p = 0.0512, two-sample Wilcoxon rank sum test).

Earlier version: We tested for differences in auditory discrimination by comparing the performance of mice at different epochs along training. First, by the end of the first week of training, all mice performed above the discrimination threshold (Fig. 1f). Although adolescent mice performed significantly fewer trials during the first week of training (Fig. 1f, vertical lines), there was no significant difference between the performance of adult and adolescent mice at the end of the first week (Fig. 1f, horizontal lines; d’ calculated from the last 100 trials). Following one week of training on the easy task, we introduced the second pair of pure tones (Fig. 1b; ‘Easy + Hard’). Immediately after introducing the hard pair, the performance of adolescent mice on the easy task dropped, while those of adults remained unaffected (Fig. 1g). During the second week of training adolescents regained their high d’ values on the easy task, such that by the end of the experiment both groups had similarly high performance on the easy task (Fig. 1h).

Revised version: We tested for differences in auditory discrimination by comparing the performance of mice at different epochs along training. First, by the end of the first week of training, all mice performed above the discrimination threshold (Fig. 1f). Although adolescent mice performed significantly fewer trials during the first week of training (Fig. 1f; vertical lines: mean # of trials), there was no significant difference between the performance of adult and adolescent mice at the end of the first week (Fig. 1f; horizontal lines: mean d’; d’ calculated from the last 100 trials). Following one week of training on the easy task, we introduced the second pair of pure tones (Fig. 1b; ‘Easy + Hard’). Immediately after introducing the hard pair, the performance of adolescent mice on the easy task dropped, while those of adults remained unaffected (Fig. 1g). During the second week of training, adolescents regained their high d’ values on the easy task, such that by the end of the experiment both groups had similarly high performance on the easy task (Fig. 1h).

Earlier version: We then compared how mice learned the hard task during the second week while training on both the easy and hard task simultaneously. While both groups now performed a similar number of trials (Fig. 1i, vertical lines), adult mice outperformed adolescent mice on the hard task (Fig. 1i, horizontal lines), suggesting that adolescents struggled with this (harder) level of discrimination. Indeed, on average, adolescent mice did not improve on the hard task during the second week of training (Fig. 1j, adolescents). Adults, on the other hand, improved their performance on the hard task during the second week (Fig. 1j). Notably, neither group could learn the hard task if it was not preceded by the easy task (Fig. S1-1a-b).

Revised version: We then compared how mice learned the hard task during the second week while training on both the easy and hard task simultaneously. While both groups now performed a similar number of trials (Fig. 1i, vertical lines: mean # of trials), adult mice outperformed adolescent mice on the hard task (Fig. 1i, horizontal lines: mean d’), suggesting that adolescents struggled with this harder level of discrimination. Indeed, on average, adolescent mice did not improve on the hard task during the second week of training (Fig. 1j, ‘adolescent’). Adults, on the other hand, improved their performance on the hard task during the second week (Fig. 1j, ‘adult’). Notably, neither group could learn the hard task if it was not preceded by the easy task (Fig. S1-1a, b).

Earlier version: Although all mice trained for exactly 14 days, each mouse performed a unique number of trials, as these were initiated by the mice spontaneously. As noted above, pre-adolescent mice performed fewer trials/day as compared to adults but, notably, only during the first week of training (mean ± SEM: adolescents, 8317 ± 712 trials, adults 13229 ± 514 trials, t-stat = -5.6314, df = 28, p = 4.9566e-06, two-sample independent t-test). Still, this difference raises the question whether the different number of trials affected our conclusions from the first week of training. Thus, to evaluate the possible connection between the number of trials and performance in our data, we carried out several analyses. First, we found that the number of trials and d’ were not correlated in either age-group during the first week of discrimination on the easy task (Pearson-r; adolescent: r = 0.1880, p = 0.5023; adult: r = -0.0730, p = 0.7959), nor on the second week of discrimination on both the easy and the hard task (Pearson-r; adolescent: r = 0. 2264, p = 0. 4171; adult: r = 0. 0194, p = 0. 9454). Second, we compared performance at three time points along the task with reference to times with shared number of trials: 1) evaluating d’ after the minimal number of trials performed by all mice during the first week of training, and 2) evaluating d’ at the mean number of trials of each group. We found no significant differences between adolescents and adults in all comparisons (Fig. S1-1c-e). Additionally, we found no significant differences between males and females, nor among experiments that had different numbers of co-housed mice in the Educage (Table 1). Taken together, despite the different absolute number of trials during the first week, we conclude that adolescents and adults learned the easy task equally quickly and effectively. The central difference between the age groups was in their ability to learn the hard task during the second week, when adults outperformed adolescents.

Revised version: Each mouse performed a unique number of trials, as these were initiated by the mice spontaneously. As noted above, pre-adolescent mice performed fewer trials/day as compared to adults but, notably, only during the first week of training (mean ± STE: adolescents, 8317 ± 712 trials, adults 13229 ± 514 trials, t-stat = −5.6314, df = 28, p = 4.9566e-06, two-sample independent t-test). Still, this difference raises the question whether the different number of trials affected our conclusions from the first week of training. Thus, to evaluate the possible connection between the number of trials and performance in our data, we carried out several analyses. First, we found that the number of trials and d’ were not correlated in either age-group during the first week of discrimination on the easy task (Pearson-r; adolescent: r = 0.1880, p = 0.5023; adult: r = −0.0730, p = 0.7959), nor on the second week of discrimination on both the easy and the hard task (Pearson-r; adolescent: r = 0. 2264, p = 0. 4171; adult: r = 0. 0194, p = 0. 9454). Second, we compared performance at three time points along the task with reference to times with shared number of trials: 1) evaluating d’ after the minimal number of trials performed by all mice during the first week of training, and 2) evaluating d’ at the mean number of trials of each group as a reference. We found no significant differences between adolescents and adults in all comparisons (Fig. S1-1c-e). Additionally, we found no significant differences between males and females, nor among experiments that had different numbers of co-housed mice in the Educage (Table 1). Taken together, despite the different absolute number of trials during the first week, we conclude that adolescents and adults learned the easy task equally quickly and effectively. The central difference between the age groups was in their ability to learn the hard task during the second week, when adults outperformed adolescents.

Earlier version: Adolescents have higher response bias and behavioral variability

Revised version: Adolescents have higher lick bias and higher behavioral variability

Earlier version: We next asked what perceptual and/or cognitive aspect of behavior is different among the groups. We plotted psychometric curves based on both learned and probe trials from the last 2000 trials of the experiment (Fig. 2a). The false alarm rate but not the hit rate of adolescent mice was significantly higher as compared to adults (Fig. 2b). To calculate the decision boundary of mice and their perceptual sensitivity we normalized the psychometric curves of each mouse using a unity-based normalization and sigmoid fitting (Fig. 2c). We found no differences in the decision boundary or slopes of the curves, suggesting that perceptual sensitivity is not different between the groups (Fig. 2c). To test whether the differences arise from a cognitive effect (i.e., lick bias), we calculated the decision threshold as the maximal lick bias, also known as the criterion bias (C). C reflects the tendency to respond in a liberal (i.e., negative C) or a conservative (i.e., positive C) manner to the sounds during the task. C was negative for both groups, yet significantly lower in adolescents, suggesting that their cognitive ability to withhold licking is inferior to that of adults (Fig. 2d).

Revised version: We next asked what perceptual and/or cognitive aspect of behavior is different among the age groups. We plotted psychometric curves based on both learned and probe trials from the last 2000 trials of the experiment (Fig. 2a). The false alarm rate but not the hit rate of adolescent mice was significantly higher as compared to adults (Fig. 2b). To calculate the decision boundary of mice and their perceptual sensitivity we normalized the psychometric curves of each mouse using a unity-based normalization and sigmoid fitting (Fig. 2c). We found no differences in the decision boundary or slopes of the curves, suggesting that perceptual sensitivity is not different between the groups (Fig. 2c). To test whether the differences arise from a cognitive effect (e.g., lick bias), we calculated the decision threshold as the maximal lick bias, also known as the criterion bias (C-bias). C-bias reflects the tendency to respond in a liberal (i.e., negative C-bias) or a conservative (i.e., positive C-bias) manner to the sounds during the task. C-bias was negative for both groups, yet significantly lower in adolescents, suggesting that their cognitive ability to withhold licking is inferior to that of adults (Fig. 2d).

Added in the revised version: Adolescent mice exhibit lower performance in the head-fixed discrimination task.

Added in the revised version: A. Experimental timeline of training followed by recordings. B. Trial structure during the recording. Solid lines indicate the tone period. Dashed lines show the reward or punishment delay (0.6 sec), and the response window (2 sec). C. Example session. Licks (grey ticks) and trial outcomes (hit = green, false alarm = yellow, miss = red and correct reject = blue) across all trials in one recording session. D. Discriminability during training sessions for the easy task (light blue) and hard task (dark blue). E. Change in d’ after the introduction of the hard task (last 100 trials of the last session of the easy task compared to last 100 trials of the first hard session; adolescent: sign-rank = 1, p = 0.0625; adult: sign-rank = 10, p = 0.9998; rank-sum = 14, p = 0.0381, two-sample Wilcoxon rank sum test). F. Expert d’ of the last 100 trials during the last training session of the easy task (rank-sum = 21, p = 0.1255, two-sample Wilcoxon rank sum test). G. Same as ‘F’, but for the hard task (rank-sum = 17, p = 0.0173, two-sample Wilcoxon rank sum test). H. Behavioral performance (average d’ of the easy and the hard task) per mouse during recording sessions for adolescents (n =13, left) and adults (n = 14, right; trials per recording: adolescent: 340.5385 ± 45.0650; adult: 431.1429 ± 30.3367; independent t-test, t-statistic = −203.7581, p = 0.1116). I. Same as ‘H’ but only for the first 148 trials. The color bar shows the p-values between the groups. J. Average cumulative licks per trial in adolescents (dashed-line) and adults (solid-line) from −200ms before tone-onset until the reward or punishment delay, 500ms after tone-offset. K. Lick latency per trial for adolescent (left) and adult (right) groups during electrophysiological recordings (LME statistics are shown in supplemental Table 1). J. Same as ‘K’ for the Lick count.

Removed in the revised version: Adolescent mice exhibit lower response inhibition and higher variability in performance over the course of learning.

Removed in the revised version: A. Schematic timeline of tone discrimination). B. Average psychometric curve of the proportion of lick responses (above the threshold) across all stimuli (probe and learned; No-Go stimuli: p = 0. 0085, Go stimuli: p = 0. 0963, Wilcoxon rank sum test, after Bonferroni correction). C. Same as ‘B’,but normalized and fitted to a sigmoid function0(proportion to the lick threshold at 10kHz, p = 0.1163; frequency at 50% proportion to the lick threshold, p = 0.5896, Wilcoxon rank sum-test, after Bonferroni correction). D. Maximal lick bias (i.e., criterion bias) per mouse during the last 2000 trials (c-bias: p = 0.0136, Wilcoxon rank sum test). E. Average psychometric curve, fitted to a sigmoid function across all stimuli (probe and learned) at different periods of the experiment — Left: first tertile-easy task (per mouse; adolescent n = 15, adult n = 15); Middle-left: last tertile – easy task; Middle-right: first tertile-hard task; Right: last tertile-hard task. (Dashed vertical =category boundary). F. Same as E. as the maximal C-bias. From left to right:(z = -0.5104, p = 0.6098, two-sample Wilcoxon rank sum test), (z = -2.6546, p = 0.0079, two-sample Wilcoxon rank sum test), (z = -4.0739, p = 4.6228e-05, two-sample Wilcoxon rank sum test), (z = -3.2768, p = 0.0011, two-sample Wilcoxon rank sum test). G. Change in lick bias (Δ c-bias) between the last 100 trials before and after the introduction of the hard task. (left, adolescent signed rank = 17, p = 0.0125; right, adult signed rank = 83, p = 0.2078, one-sample Wilcoxon signed rank test; Δ c-bias between adolescents and adults, z = -2.9035; p = 0.0037, two-sample Wilcoxon rank sum test). H. Coefficient of variation (CV) of the d’ of the easy task before the introduction of the hard task. (z = -0.1659, p = 0.8682, two-sample Wilcoxon rank sum test) I. Same as ‘H’ after the introduction of the hard task (z = 2.1569, p = 0.0310, two-sample Wilcoxon rank sum test). J. Same as ‘H’ but for d’ of the hard task (z = 2.5302, p = 0.0114, two-sample Wilcoxon rank sum test).

Earlier version: Given that adolescents and adults show different dynamics in their performance along the task (Fig. 1), we tested whether the differences we found in the lick bias are consistent at different times along the task. To normalize for different numbers of trials in each mouse, we divided their training episodes into the first and last tertiles during the first week and second week of training (Fig.2e). The maximal lick bias was significantly higher in adolescents in all but the very first training episode, when mice learned the procedure (Fig. 2e-f). These data highlight the general cognitive difficulty of adolescents to withhold licking during the task. The cognitive sensitivity of adolescents was further evident when their performance dropped right after the task was switched from ‘easy only’ to ‘easy + hard’ (Fig. 1f). The lick bias of adolescent mice became significantly stronger after the introduction of the hard task (Fig. 2g). In contrast, the lick bias of adult mice remained stable (Fig. 2g). These data further suggest that this cognitive trait (i.e., to withhold licking) has not yet matured in the young mice. To test if the lick bias was stimulus-related or a result of general impulsivity, we also compared the inter-trial interval (ITI) after different trial outcomes, and observed similar ITIs across groups, except for false alarm trials, where adolescents initiated trials significantly faster (Fig. S2-1). This suggests that adolescent mice exhibit more impulsive responses following punishments, contributing to their immature lick bias.

Revised version: Given that adolescents and adults show different dynamics in their performance along the task (Fig. 1), we tested whether the differences we found in the lick bias are consistent at different times along the task. To normalize for different numbers of trials in each mouse, we divided their training episodes into the first and last tertiles during the first week and second week of training separately (Fig. 2e). The maximal lick bias was significantly higher in adolescents in all but the very first training episode, when mice learned the procedure (Fig. 2f). These data highlight the general cognitive difficulty of adolescents to withhold licking during the task.

Added in the revised version: The cognitive sensitivity of adolescents was further evident when their performance dropped right after the task was switched from ‘easy only’ to ‘easy + hard’ (Fig. 1f). The lick bias of adolescent mice became significantly stronger after the introduction of the hard task (Fig. 2g). In contrast, the lick bias of adult mice remained stable (Fig. 2g). These data further suggest that this cognitive trait (i.e., to withhold licking) has not yet matured in the young mice. To test if the lick bias was stimulus-related or a result of general impulsivity, we also compared the inter-trial interval (ITI) after different trial outcomes, and observed similar ITIs across groups, except for false alarm trials, where adolescents initiated the trials significantly faster (Fig. S2-1). This, too, suggests that adolescent mice exhibit more impulsive responses following punishments (or alternatively are less impacted by punishments), contributing to their immature lick bias.

Earlier version: Prior studies have shown that a dominant feature of adolescent behavior is that it is more variable compared to adults[4]. To examine behavioral variability in our data, we calculated the coefficient of variation (CV) of the discriminability (d’) for each mouse across training. While there was no significant difference in d’ variability when training on the easy task during the first week of training (Fig. 2h), adult mice showed significantly lower variability during the second week of training while they trained on both the easy and hard tasks (Fig. 2i, j). Thus, adolescent learning and behavior are characterized by both lower response inhibition and higher variability, particularly during the second week when the task involved more challenging discrimination.

Revised version: Prior studies have shown that a dominant feature of adolescent behavior is that it is more variable compared to adults[4]. To examine behavioral variability in our data, we calculated the coefficient of variation (CV) of the discriminability (d’) for each mouse across training. While there was no significant difference in d’ variability when training on the easy task during the first week of training (Fig. 2h), adult mice showed significantly lower variability during the second week of training while training on the ‘easy + hard’ versions of the task (Fig. 2i,j). Thus, adolescent learning and behavior are characterized by both lower response inhibition and higher variability, particularly during the second week when the task involved more challenging discrimination for the younger mice.

Adult mice outperform adolescent mice on a head-fixed discrimination task

Earlier version: To enable electrophysiological recordings during expert task performance, we trained adolescent (n = 5) and adult (n = 6) mice on a head-fixed learning paradigm, using a similar protocol as in the Educage (Fig. 3a). In the head-fixed protocol, water access was limited to training sessions. Head-fixed mice were trained to lick after a 100ms tone and were rewarded (with water) or punished (by a 2 sec white noise) after a 600ms delay (Fig. 3b). Each session was concluded after the mice became satiated and stopped licking. Water supply was limited to 0.0125 ml per day per gram of body weight, and mice that did not consume this amount were compensated after training (see Methods) (Fig. 3c shows an example session from one mouse). The number of trials per training session was significantly lower in adolescents (mean ± SEM: adolescents: 410 ± 35, adults: 580 ± 37 t-statistic = -27.9855, p = 0.0012, independent two-sample t-test). However, the number of sessions to reach the behavioral threshold (d’ >= 1 in the easy task) was not significantly different between adolescents and adults (mean ± SEM: adolescents: 3.7 ± 0.8, adults: 3.8 ± 0.7; t-statistic = - 0.8531, p = 0.1971, independent two-sample t-test, Fig. 3d). While the number of trials performed within each session is likely related to physiological factors, such as smaller body weight and faster satiation, they did not impact learning capabilities. Similar to learning in the Educage, adolescent performance (but not that of adults) decreased after we introduced the hard task (Fig. 3e). Another similarity between the head-fixed and Educage versions of the task was that we found no significant differences between adolescents and adults in the performance of the easy task (Fig. 3f), and that adults outperformed adolescents on the hard sound discrimination (Fig. 3g). Also similar to the behavior in the Educage, lick rates of adolescent mice were generally higher (Fig. S3-1a), and adolescents exhibited a stronger lick bias (Fig. S3-1b). In contrast, in the head-fixed task we found no evidence in impulsiveness; there were no differences between the groups in the licking patterns during the inter-trial intervals (Fig. S3-1c, d).

Revised version: To enable electrophysiological recordings during expert task performance, we trained adolescent (n = 5) and adult (n = 6) mice on a head-fixed learning paradigm, using a similar protocol as in the Educage (Fig. 3a). In the head-fixed protocol, water access was limited to training sessions. Head-fixed mice were trained to lick after a 100ms tone and were rewarded (with water) or punished (by a 2 sec white noise) only after a 600ms delay (Fig. 3b). Each session was concluded after the mice became satiated and stopped licking. Water supply was limited to 0.0125 ml per day per gram of body weight, and mice that did not consume this amount were compensated after training (see Methods; Fig. 3c shows an example session from one mouse). The number of trials per training session was significantly lower in adolescents (mean ± STE: adolescents: 410 ± 35, adults: 580 ± 37 t-statistic = −27.9855, p = 0.0012, independent two-sample t-test). However, the number of sessions to reach the behavioral threshold (d’ >= 1 in the easy task) was not significantly different between adolescents and adults (mean ± STE: adolescents: 3.7 ± 0.8, adults: 3.8 ± 0.7; t-statistic = - 0.8531, p = 0.1971, independent two-sample t-test, Fig. 3d shows all learning curves). Similar to learning in the Educage, adolescent performance (but not that of adults) decreased after we introduced the hard task (Fig. 3e). Another similarity between the head-fixed and Educage versions of the task was that we found no significant differences between adolescents and adults in the performance of the easy task (Fig. 3f), and that adults outperformed adolescents on the hard discrimination (Fig. 3g). Also similar to the behavior in the Educage, lick responses of adolescent mice were generally higher (Fig. S3-1a), and adolescents exhibited a stronger lick bias (Fig. S3-1b). In line with this result, in the head-fixed task we found higher impulsiveness in the proportion and the number of licks during the inter-trial intervals after FA trials in adolescents (Fig. S3-1c, d).

Adolescent mice exhibit lower performance in the head-fixed discrimination task.

Earlier version: A. Experimental timeline of training followed by recordings. Created with BioRender.com. B. Trial structure during the recording. Solid lines indicate the tone period. Dashed lines show the reward or punishment delay (0.6 sec), and the response window (2 sec). C. Example session. Licks (grey ticks) and trial outcomes (hit = green, false alarm = yellow, miss = red and correct reject = blue) across all trials in one recording session. D. Discriminability during training sessions for the easy task (light blue) and hard task (dark blue). E. Change in d’ after the introduction of the hard task (last 100 trials of the last session of the easy task compared to last 100 trials of the first hard session; rank-sum = 14; p = 0.0381, two-sample Wilcoxon rank sum test). F. Expert d’ of the last 100 trials during the last training session of the easy task (rank-sum = 21; p = 0.1255, two-sample Wilcoxon rank sum test). G. Same as ‘F’, but for the hard task (rank-sum = 17; p = 0.0173, two-sample Wilcoxon rank sum test). H. Behavioral performance (average d’ of the easy and the hard task) per mouse during recording sessions for adolescents (n =13, left) and adults (n= 14, right) (trials per recording: adolescent: 340.5385 ± 45.0650; adult: 431.1429 ± 30.3367; independent t-test, t-statistic = -203.7581, p = 0.1116). I. Same as ‘H’ but only for the first 148 trials. The color bar shows the p-values between the groups. J. Average cumulative licks per trial in adolescents (dashed-line) and adults (solid-line) from -200ms before tone - onset until the reward or punishment delay, 500ms after tone-offset. K. Lick latency per trial for adolescent (left) and adult (right) groups during electrophysiological recordings (LME statistics are shown in supplemental Table 1). J. Same as ‘K’ for the Lick count.

Revised version: A. Experimental timeline of training followed by recordings. Created with BioRender.com. B. Trial structure during the recording. Solid lines indicate the tone period. Dashed lines show the reward or punishment delay (0.6 sec), and the response window (2 sec). C. Example session. Licks (grey ticks) and trial outcomes (hit = green, false alarm = yellow, miss = red and correct reject = blue) across all trials in one recording session. D. Discriminability during training sessions for the easy task (light blue) and hard task (dark blue). E. Change in d’ after the introduction of the hard task (last 100 trials of the last session of the easy task compared to last 100 trials of the first hard session; adolescent: sign-rank = 1, p = 0.0625; adult: sign-rank = 10, p = 0.9998; rank-sum = 14; p = 0.0381, two-sample Wilcoxon rank sum test). F. Expert d’ of the last 100 trials during the last training session of the easy task (rank-sum = 21; p = 0.1255, two-sample Wilcoxon rank sum test). G. Same as ‘F’, but for the hard task (rank-sum = 17; p = 0.0173, two-sample Wilcoxon rank sum test). H. Behavioral performance (average d’ of the easy and the hard task) per mouse during recording sessions for adolescents (n =13, left) and adults (n= 14, right; trials per recording: adolescent: 340.5385 ± 45.0650; adult: 431.1429 ± 30.3367; independent t-test, t-statistic = −203.7581, p = 0.1116). I. Same as ‘H’ but only for the first 148 trials. The color bar shows the p-values between the groups. J. Average cumulative licks per trial in adolescents (dashed-line) and adults (solid-line) from −200ms before tone-onset until the reward or punishment delay, 500ms after tone-offset. K. Lick latency per trial for adolescent (left) and adult (right) groups during electrophysiological recordings (LME statistics are shown in supplemental Table 1). J. Same as ‘K’ for the Lick count.

Earlier version: To test whether ACx was necessary for expert discrimination in this task, we performed a causal experiment in adult mice. Adult mice (n=3) were injected bilaterally with AAV5-CAMKII-GtACR2-FRED_kv_2.1 into the ACx, and optical fibers were implanted over the injection sites (Fig. S3-2a,c; Fig. S3-3). This viral construct enables optogenetic silencing of excitatory neurons via GtACR2 (Guillardia theta Anion Channelrhodopsin 2), under the control of the CaMKII promoter. It includes a RED Fluorescent protein for expression verification (FRED) and a soma-targeting motif derived from the Kv2.1 potassium channel to enhance membrane localization. Mice were trained identically on the head-fixed Go/No-Go paradigm as described above. At the end of training, we attached a light source to the implanted fibers to allow optogenetic suppression during several sessions (Fig. S3-2a). In these sessions, light was applied in 50% of trials in a pseudo-random fashion, with light ON trials starting from -50ms prior to tone onset up to 50ms after tone offset (Fig. S3-2b). A control group of mice (n= 3) were injected with AAV9-CaMKII-dTomato and went through the exact same procedure (Fig. S3-2a,c). We compared the lick rates for both Go and No-Go stimuli under light ON and light OFF conditions. We found that lick rates were strongly affected in the GtACR2-injected mice (n = 13 sessions), but not in controls (n = 8 sessions) (Fig. S3-2d). Light ON conditions only affected the lick rate of the inhibited trial and not the follow-up trials (Fig. S3-2e). Discriminability (d’) decreased significantly with optogenetic suppression for both easy and hard tasks, although mice still performed better in the easy task under both conditions (Fig. S3-2f). These results confirm findings by others in similar (though not identical) tasks [19, 20], suggesting that the ACx is necessary for executing the behavior during the expert stage of this task (but see Discussion).

Revised version: To test whether ACx was necessary for expert discrimination in this task, we performed a causal experiment in adult mice. Adult mice (n=3) were injected bilaterally with AAV5-CAMKII-GtACR2-FRED_kv_2.1 into the ACx, and optical fibers were implanted over the injection sites (Fig. S3-2a, c; Fig. S3-3). This viral construct enables optogenetic silencing of excitatory neurons via GtACR2 (Guillardia theta Anion Channelrhodopsin 2), under the control of the CaMKII promoter. It includes a red fluorescent protein for expression verification (FRED) and a soma-targeting motif derived from the Kv2.1 potassium channel to enhance membrane localization.

Added in the revised version: Mice were trained identically on the head-fixed Go/No-Go paradigm as described above. At the end of training, we attached a light source to the implanted fibers to allow optogenetic suppression during several sessions (Fig. S3-2a). In these sessions, light was applied in 50% of trials in a pseudo-random fashion, with light-on trials starting from −50ms prior to tone onset up to 50ms after tone offset (Fig. S3-2b). A control group of mice (n= 3) were injected with AAV9-CaMKII-dTomato and went through the exact same procedure (Fig. S3-2a, c). We compared the lick responses for both Go and No-Go stimuli under light-on and light-off conditions. We found that lick responses were strongly affected in the GtACR2-injected mice (n = 13 sessions), but not in controls (n = 8 sessions) (Fig. S3-2d). Light-on conditions only affected the lick responses of the inhibited trial and not the light-off trials (Fig. S3-2e). Discriminability (d’) decreased significantly with optogenetic suppression for both easy and hard tasks, although all mice still performed better in the easy task under both conditions (Fig. S3-2f). These results confirm findings by others in similar (though not identical) tasks [19, 20], suggesting that the ACx is necessary for executing the behavior during the expert stage of this task (but see Discussion for limitations of this experiment).

Earlier version: To measure neural responses during behavioral performance, head-fixed expert mice underwent multiple recording sessions targeting the ACx (adults, n = 14 recording sessions from 6 mice, age: P77-P82; adolescents, n = 13 recordings from 5 mice, age: P37-P42). During the recordings, the performance of all mice on the easy task was at least d’>1, and the d’ on the hard task was smaller for both groups (often smaller than 1) (Fig. S3-4a). Performance was not significantly different across the multiple recording sessions in each mouse (Fig. S3-4b). Behavioral performance during recordings was not significantly different to performance in the Educage (Fig. S3-4c). Figure 3h shows the evolution of discriminability (d’) during all recordings for adolescents and adults. The performance was heterogeneous across mice as well as within recording-sessions. However, the total number of trials performed was not significantly different (trials per recording: adolescent: 340.5 ± 45; adult: 431.1 ± 30.3; independent t-test, t-statistic = -203.7581, p = 0.1116).

Revised version: To measure neural responses during behavioral performance, head-fixed expert mice underwent multiple recording sessions targeting the ACx (adults, n = 14 recording sessions from 6 mice, age: P77-P82; adolescents, n = 13 recordings from 5 mice, age: P37-P42). During the recordings, the performance of all mice on the easy task was at least d’>1, and the d’ on the hard task was smaller for both groups (often smaller than 1; Fig. S3-4a). Performance was not significantly different across the multiple recording sessions in each mouse (Fig. S3-4b). Behavioral performance during recordings was not significantly different to performance in the Educage (Fig. S3-4c). Figure 3h shows the evolution of behavioral discriminability (d’) during all recordings for adolescents and adults. The performance was heterogeneous across mice as well as within recording-sessions. However, the total number of trials performed was not significantly different between age groups (trials per recording: adolescent: 340.5 ± 45; adult: 431.1 ± 30.3; independent t-test, t-statistic = −203.7581, p = 0.1116).

Earlier version: We found additional differences between the behavior of adults and adolescents in the head restrained configuration. The average adolescent discriminability was lower at the beginning of the session (initial 78 trials) but then rapidly leveled out (Fig. 3i). Adolescent mice had shorter lick latencies (Fig. 3j, k), and higher lick counts (Fig. 3j, l; Supplementary Table 1). While the lick latencies were independent of the discriminability (d’), we found a significant interaction effect between lick count and the d’ (Supplementary Table 1). Together, these behavioral differences suggest that the weaker response inhibition (here, expressed as lick latency) and higher reward anticipation (here, expressed as lick count) may contribute to the weaker performance of adolescents during the beginning of the recording.

Revised version: We found additional differences between the behavior of adults and adolescents in the head restrained configuration. The average adolescent discriminability was lower at the beginning of the session (initial 78 trials) but then rapidly leveled out (Fig. 3i). Figure 3j shows the cumulative lick curves of adults and adolescents. Adolescent mice had shorter lick latencies (Fig. 3k), and higher lick counts (Fig. 3l; Supplementary Table 1). While the lick latencies were independent of the discriminability (d’), we found a significant interaction effect between lick count and the d’ (Supplementary Table 1). Together, these behavioral differences suggest that the weaker response inhibition (here, expressed as lick latency) and higher reward anticipation (here, expressed as lick count) may contribute to the weaker performance of adolescents during the beginning of the recording session.

Neuronal representations of stimulus- and choice-related activity are immature in adolescents

Earlier version: To compare the activity of neurons in adolescent and adult mice, we recorded spiking activity from the ACx using the high-density Neuropixels-1 probes in expert mice engaged in the task. We targeted the ACx by inserting the probe in a diagonal angle, traversing four auditory regions: Dorsal Auditory Cortex (AUDd), Primary Auditory Cortex (AUDp), ventral auditory cortex (AUDv), and Temporal Association Cortex (TEa) (Fig. 4a left, Fig. S4-1a-b show an example voltage trace and one example recording). We recorded multiple times from each mouse and used DiI- or DiO-coatings for different penetrations to verify probe trajectories postmortem. We used the 3D-allen CCF-slice reconstruction (github.com/cortex-lab/allenCCF), for high resolution anatomical registration (Fig.4a, b; Fig. S4-1c, d).

Revised version: To compare the activity of neurons in adolescent and adult mice, we recorded spiking activity from the ACx using the high-density Neuropixels-1 probes in expert mice engaged in the task. We targeted the ACx by inserting the probe in a diagonal angle, traversing four auditory regions: Dorsal Auditory Cortex (AUDd), Primary Auditory Cortex (AUDp), ventral auditory cortex (AUDv), and Temporal Association Cortex (Tea; Fig. 4a; Fig. S4-1a,b, show an example voltage trace and one example recording). We recorded multiple times from each mouse and used DiI-or DiO-coatings for different penetrations to verify probe trajectories postmortem. We used the 3D-allen CCF-slice reconstruction platform (github.com/cortex-lab/allenCCF), for high resolution anatomical registration (Fig.4b; Fig. S4-1c, d).

ACx neurons in adolescents exhibit lower discriminability in stimulus- and choice-related activity.

Earlier version: A. Recordings in ACx when the mouse is engaged in the task, using Neuropixels-1 probes. Left: Recordings were performed in AUDd, AUDp, AUDv, and TEa. Right: Fluorescent micrograph of a coronal brain slice showing the probe tracks of three recordings (red = DiI, yellow = DiO). Created with BioRender.com. B. Top: 3D-Reconstruction of recording sites in adolescents (n = 13; grey) and adults (n = 14; black). Bottom: distribution of the spike-depth of all excitatory tone-responsive L5/6 neurons in adolescents (n = 455; grey) and adults (n = 607; black). C. Normalized PSTH and lick rate (LR) from -200ms to +600ms after tone-onset in adolescents (grey) and adults (black). D. Spiking activity from one example neuron sorted by trial outcome (hit, miss, false alarm, correct reject). Top: PSTH per trial outcome. Bottom: Heat map of the FR sorted pertrial outcome. E. Discriminability values (AUC) over time (from - 200ms to 600ms after tone onset) for one example neuron (same neuron as in ‘D’). AUC values are shown for stimulus related activity (left: easy task, middle: hard task) and choice-related activity (right). Shuffled distribution in all curves is shown in grey. F. Same as ‘E’ for all neurons. The curves are average (+-SEM) neuronal discriminability of adult neurons (solid line) and adolescent neurons (dashed line), for easy (adolescent neurons = 190, mice = 4, recordings = 7; adult n = 358, mice = 4, recordings = 8; left) and hard stimulus-related activity (adolescent n = 429; adult n = 562 , mice = 5, recordings = 9; middle), and choice-related activity (adolescent n = 429; adult n = 562 , mice = 5, recordings = 9; right). G. 3D plots of the onset-latency of discriminability (ms), duration of discriminability (ms), and maximal discriminability (AUC) of all neurons that showed significant discriminability. Left: easy task (adolescent neurons = 178 (93%), mice = 4, recordings = 6; adult n = 346 (97%), mice = 4, recordings = 8; left); Center: hard task(adolescent neurons = 399 (93%), mice = 5, recordings = 10; adult n = 544 (97%), mice = 6, recordings = 12; middle); Right: choice-related activity (adolescent neurons = 181 (95%), mice = 4, recordings = 9; adult n = 339 (95%), mice = 4, recordings = 7; right).

Revised version: A. Recordings in ACx when the mouse is engaged in the task, using Neuropixels-1 probes. Left: Recordings were performed in AUDd, AUDp, AUDv, and TEa. Right: Fluorescent micrograph of a coronal brain slice showing the probe tracks of three recordings (red = DiI, yellow = DiO). Created with BioRender.com. B. Top: 3D-Reconstruction of recording sites in adolescents (n = 13; grey) and adults (n = 14; black). Bottom: distribution of the spike-depth of all excitatory tone-responsive L5/6 neurons in adolescents (n = 455; grey) and adults (n = 607; black). C. Normalized PSTH (FR in Hz) and lick-rate (LR in Hz) from −200ms to +600ms after tone-onset in adolescents (grey) and adults (black). D. Spiking activity from one example neuron sorted by trial outcome (hit, miss, false alarm, correct reject). Top: PSTH per trial outcome. Bottom: Heat map of the FR sorted per trial outcome. E. Discriminability values (AUC) over time (from −200ms to 600ms after tone onset) for one example neuron (same neuron as in ‘D’). AUC values are shown for stimulus related activity (left: easy task, middle: hard task) and choice-related activity (right). Shuffled distribution in all curves is shown in grey. F. Same as ‘E’ for all neurons. The curves are average (+-STE) neuronal discriminability of adult neurons (solid line) and adolescent neurons (dashed line), for easy (adolescent neurons = 190, mice = 4, recordings = 7; adult n = 358, mice = 4, recordings = 8; left) and hard stimulus-related activity (adolescent n = 429; adult n = 562, mice = 5, recordings = 9; middle), and choice-related activity (adolescent n = 429; adult n = 562, mice = 5, recordings = 9; right). G. 3D plots of the onset-latency of discriminability (ms), duration of discriminability (ms), and maximal discriminability (AUC) of all neurons that showed significant discriminability. Left: easy task (adolescent neurons = 178 (93%), mice = 4, recordings = 6; adult n = 346 (97%), mice = 4, recordings = 8; left); Center: hard task (adolescent neurons = 399 (93%), mice = 5, recordings = 10; adult n = 544 (97%), mice = 6, recordings = 12; middle); Right: choice-related activity (adolescent neurons = 181 (95%), mice = 4, recordings = 9; adult n = 339 (95%), mice = 4, recordings = 7; right).

Earlier version: We limited our single neuron analysis to well-isolated single units from infragranular layer 5 and layer 6 (recordings were restricted to these layers due to the probe angle), as two main projection layers of the cortex, which are key nodes for assessing cortical outputs [21]. Recorded neurons from other layers and areas were excluded from the analysis. Together, we collected data from 1145 single neurons in 13 recordings of 5 adolescent mice and 1267 single neurons in 14 recordings of 6 adult mice (Supplementary Table 2; an overview of every dataset can be found in Supplementary Table 7). We further restricted our analysis to single neurons excited by sounds during the first 150ms after tone onset, to exclude any contributions of motor actions due to licking (Fig. 4c; Fig. S4-1e). Lick rate- and Firing rate-PSTHs were not correlated during the first 150ms (adolescents r = -0.1515, p = 0.9878; adults r = -0.0638, p = 0.9942). In total, we analyzed n = 463 neurons from adolescent mice and n = 599 from adults.

Revised version: We limited our single neuron analysis to well-isolated single units from infragranular layer 5 and layer 6 (recordings were largely restricted to these layers due to the probe angle), as two main projection layers of the cortex, which are key nodes for assessing cortical outputs [21]. Recorded neurons from other layers and areas were excluded from the analysis. Together, we collected data from 1145 single neurons in 13 recordings of 5 adolescent mice and 1267 single neurons in 14 recordings of 6 adult mice (Supplementary Table 2; an overview of every dataset per figure is detailed in Supplementary Table 9). We further restricted our analysis to single neurons excited by sounds during the first 150ms after tone onset. Figure 4c shows the normalized firing rate (FR) compared to the normalized lick rate (LR) from −200ms to 600ms after the tone onset, before the end of the reinforcement delay (Fig. S4-1e shows the FR in Hz per auditory sub-region). FR-PSTHs and LR-PSTHs were not correlated during the first 150ms (adolescents r = −0.1515, p = 0.9878; adults r = −0.0638, p = 0.9942). In total, we analyzed n = 463 neurons from adolescent mice and n = 599 from adults.

Earlier version: In Table 2, we present the number of neurons in our dataset, categorized by region. A detailed overview of the differences in firing properties per auditory region is also provided in Figure S4-2 and Supplementary Table 3. Consistent with previous work, we found that firing properties such as minimal latency differed between auditory regions in both age groups, with AUDp and TEa being most distinct (Fig. S4-2, Supplementary Table 3) [22].

Revised version: In Supplementary Table 2, we present the number of neurons in our dataset, categorized by region. A detailed overview of the differences in firing properties per auditory region is also provided in Supplementary Table 3 and Supplementary Table 4.

Added in the revised version: The differences in firing properties between adolescent and adult neurons varied across auditory subregions. AUDp showed differences in all firing properties that we analyzed. In AUDd, adolescent neurons had greater latency to peak, full-width half-maximum, minimal latency, and lower trial-responsiveness. AUDv also showed greater minimal latency and lower lifetime sparseness. TEa neurons of adolescents exhibited greater minimal latency and higher spontaneous FR. While these differences may partially reflect the varying number of neurons recorded per subregion and different functional attributes of each region, we observed similar trends in Cohen’s D values (effect sizes) across all four areas, supporting a generally consistent group differences in activity patterns (Supplementary Table 3). Consistent with previous work, we also found that firing properties such as minimal latency differed between auditory regions in both age groups (Supplementary Table 4) [22].

Removed in the revised version: Notably, the effect size of the differences in baseline firing rates (FR) between groups was larger in AUDd, AUDp, and AUDv. However, the evoked FR was smaller in all four areas. The coefficient of variance of the FR was larger in AUDd, AUDp, and AUDv. Additionally, the latency to peak was larger in AUDd and AUDp, while the full-width half-maximum of the FR was also larger in these two areas. The minimal latency was larger in all four areas, and the number of responsive trials was smaller in all four areas. Finally, the lifetime sparseness (see Methods) was smaller in AUDp and AUDv between adolescent and adult neurons (Supplementary Table 3).

Earlier version: Next, to study how adolescents and adults encode task performance we divided neuronal responses by trial outcome (Fig. 4d). Specifically, we assessed neuronal discriminability between stimulus-related and choice-related activity by calculating the area under the curve (AUC) from a receiver operating curve (ROC) (Fig. 4e). Stimulus-related activity was calculated as the difference between hit and false alarm trials in the easy and the hard task separately. Choice-related activity was calculated as the average difference between false-alarm and correct reject trials (miss trials were excluded since we had insufficient trial numbers of this outcome). Figure 4e shows the average AUC values over time from an example neuron (same neurons shown in Fig. 4D), and figure 4f for all neurons. The vast majority of neurons successfully discriminated (AUC above shuffled data) either or both stimulus- and choice-related activity (adolescent: easy stimuli = 93%, hard stimuli = 93%, choice= 95%; adult: easy stimuli = 97%, hard stimuli = 97%, choice= 95%). The onset-latency of discriminability was significantly slower, maximal discriminability significantly weaker, and the duration of discriminability significantly shorter in adolescent neurons compared to adult neurons (Table 2; plotted in Fig. 4g). We found no differences between auditory regions (Fig. S4-3). Thus, auditory processing of stimulus- and choice-related activity in the adolescent ACx has not yet reached full maturity.

Revised version: To study how adolescents and adults encode task performance we divided neuronal responses by trial outcome (Fig. 4d shows one example). Specifically, we assessed neuronal discriminability between stimulus-related and choice-related activity by calculating the area under the curve (AUC) from a receiver operating curve (ROC). Specifically, stimulus-related activity was calculated as the difference between hit and false alarm trials in the easy and the hard task separately (Fig. 4e, left). Choice-related activity was calculated as the average difference between false-alarm and correct reject trials (Fig. 4e, left). Miss trials were excluded since we had insufficient trial numbers of this outcome. Figure 4e (right) shows the average AUC values over time from an example neuron (same neuron shown in Fig. 4D), and figure 4f shows the average calculated from all neurons. The vast majority of neurons successfully discriminated (AUC above shuffled data) either or both stimulus- and choice-related activity (adolescent: easy stimuli = 93%, hard stimuli = 93%, choice= 95%; adult: easy stimuli = 97%, hard stimuli = 97%, choice= 95%). The onset-latency of discriminability was significantly slower, maximal discriminability significantly weaker, and the duration of discriminability significantly shorter in adolescent neurons compared to adult neurons (discriminability traces are plotted in Fig. 4f, and all values from individual neurons and the 3 metrics are plotted together in Fig. 4g; see Table 2 for statistics). The four auditory subregions were not significantly different in all 3 measured metrics (Fig. S4-2). Thus, auditory processing of stimulus- and choice-related activity in the adolescent ACx has not yet reached full maturity.

Earlier version: Population decoding of the hard task in adults outperforms adolescents

Revised version: Adults exhibit better population decoding of the difficult task compared to adolescents

Earlier version: While individual neurons exhibited weaker discriminability in adolescents than in adults, at a population level these effects could saturate out. We therefore tested the ability of recorded populations to decode trial outcomes. We used a linear discriminant analysis (LDA) decoder to quantify hit versus correct reject trial outcomes. This comparison extracts information of the combination of stimulus and choice (i.e. task-related differences). We decoded separately in easy trials and hard trials, based on population activity during the first 200 milliseconds. For the easy task there was no significant difference between decoding accuracy in adolescents and adults (Fig. 5a). However, decoding accuracy was significantly higher in adults compared to adolescents in hard trials (Fig. 5a). For both adults and adolescents, decoding accuracy was significantly lower in hard trials compared to easy trials (Fig. 5a). We performed finer timescale decoding to quantify the onset latency of differences between the activity for hit vs. correct reject trials, defined as the time in which the decoding accuracy first crosses three times the standard deviation away from the baseline (see Methods). We found that the latency was shorter in adults in both the easy and hard trials (Fig. 5b). In adolescent animals the latency was larger in hard trials (Fig. 5b) but not in adult animals (Fig. 5b). We then extended this analysis over a temporal window from -0.5s to10s relative to tone onset (Fig. 5c). In the 200ms window after the response period (2.0-2.2 s), there is no longer a significant difference in decoding accuracy between adult and adolescent mice as well as between the two tasks (Fig. S5-1).

Revised version: While individual neurons exhibited weaker discriminability in adolescents than in adults, at a population level these effects could saturate out. We therefore tested the ability of recorded populations to decode trial outcomes. We used a linear discriminant analysis (LDA) decoder to quantify hit versus correct reject trial outcomes. We applied stringent inclusion criteria for the number of simultaneously recorded neurons (see Methods), which restricted our analysis to hit versus correct rejection trials with sufficient trial counts for statistical reliability. We note that this comparison extracts information of the combination of stimulus and choice information indiscriminately (i.e. collectively it includes the full ‘task-related’ differences). We decoded separately in easy trials and hard trials, based on population activity during the first 200 milliseconds. For the easy task there was no significant difference between decoding accuracy in adolescents and adults (Fig. 5a). However, decoding accuracy was significantly higher in adults compared to adolescents in hard trials (Fig. 5a). For both adults and adolescents, decoding accuracy was significantly lower in hard trials compared to easy trials (Fig. 5a). We performed finer timescale decoding to quantify the onset latency of differences between the activity for hit vs. correct reject trials, defined as the time in which the decoding accuracy first crosses three times the standard deviation away from the baseline (see Methods). We found that the latency was shorter in adults in both the easy and hard trials (Fig. 5b). In adolescent animals the latency was larger in hard trials (Fig. 5b) but not in adult animals (Fig. 5b). We then extended this analysis over a temporal window from −0.5s to10s relative to tone onset (Fig. 5c). In the 200ms window after the response period (2.0-2.2 s), there is no longer a significant difference in decoding accuracy between adult and adolescent mice as well as between the two task difficulties (Fig. S5-1).

Decoding in adult neuronal populations outperforms decoding in adolescents.

A. Decoding accuracy for the first 200ms across all recordings in both adults (black) and adolescents (grey) for the easy task (adolescents compared to adults, p = 0.5000, Student’s t-test) and the hard task (adolescents compared to adults, p = 0.0300, Student’s t-test). Decoding is better in the easy task for both age groups (adults: p = 0.0030; adolescents: p = 0.01, paired t-test). B. Decoding latency for all recordings in the easy task (p = 0.0200, Student’s t-test) and the hard task (p = 0.0030, Student’s t-test), as well as compared between age groups (easy task, p = 0.05400, pared t-test; hard task: p = 0.0100). C. Decoding accuracy over a time window from −0.5s to 10s (the response window highlighted in the grey) for the easy task (left) and the hard task (right). D. LDA separation for easy and hard tasks. Lines represent robust linear regression fits without intercept (Huber loss; robust linear regression, p = 0.0001) E. Single trial variance for easy and hard tasks in adolescent and adult recordings (adults: p = 0.0040; adolescents: p = 0.0300, paired t-test; easy task: p = 0.4500; hard task: p=0.4100, Student’s t-test). F. Visualization of population representations for the stimuli in easy and hard tasks. Dotted lines indicate decoding dimensions, and ellipses represent the covariance of the representations.

Earlier version: To further quantify the different decoding performance between two tasks, we plotted the fisher separation metric ( see Methods) for each level of difficulty. The separation ratio of hard over easy tasks was significantly higher in adult mice than in adolescent mice (robust linear regression, see Methods). Separation could have changed due to alterations in the mean or the dispersion around the mean. Testing the dispersion, we found that the variance around the means was not equal across easy and hard trials but rather significantly increased for the hard trials (Fig. 5d-f). However, there was no significant difference in variance between adolescent and adult mice within the same task (Fig. 5d-f).

Revised version: To further quantify the different decoding performance between two tasks difficulties, we plotted the fisher separation metric (, see Methods) for each level of difficulty. While decoding accuracy reflects how well a classifier can distinguish between trial types, the Fisher separation metric specifically quantifies the ratio of between-category variance to within-category variance. The separation ratio of hard over easy tasks was significantly higher in adult mice than in adolescent mice (robust linear regression, see Methods). Separation could have changed due to alterations in the mean or the dispersion around the mean. Testing the dispersion, we found that the variance around the means was not equal across easy and hard trials but rather significantly increased for the hard trials (Fig. 5d-f). However, there was no significant difference in variance between adolescent and adult mice within the same task (Fig. 5d-f).

Earlier version: Effects of age and learning on cortical plasticity in mice engaged in the task

Revised version: Both age and learning affect cortical plasticity in mice engaged in the task

Earlier version: We next studied how learning contributes to cortical plasticity in the different age groups. To separate between age- and learning-related effects we recorded from the ACx of two new groups of mice (adolescents and adults) that are novice on the task (Fig. 6a, P38, recordings= 6; P78, recordings =6, n = 3 mice per group). We collected 130 tone-responsive (by excitation) neurons in novice adolescents and 186 tone-responsive (by excitation) neurons in novice adults (Fig. 6b-c; Supplementary Table 2). Notably, novice mice did not discriminate between sounds; instead, they licked indiscriminately to the tones presented during the task (Fig. S61a-b).

Revised version: We next studied how learning contributes to cortical plasticity in the different age groups. To separate between age- and learning-related effects, we recorded from the ACx of two new groups of mice (adolescents and adults that are age-matched to expert mice) that are novice on the task (Fig. 6a, novice P37-P42, recordings= 6; novice P77-P82, recordings =6, n = 3 mice per group). We collected data from 130 tone-responsive (by excitation) neurons in novice adolescents and 186 tone-responsive (by excitation) neurons in novice adults (Fig. 6b; Supplementary Table 2). Figure 6c shows the normalized average FR PSTH together with the LR PSTH for novice mice. Notably, novice mice did not discriminate between sounds (Fig. S6-1a, b). Similar to expert mice (Supplementary Table 3), we found diverse firing properties across auditory subregions in novice mice (Supplementary Table 5). Also similar to expert mice, several basic firing properties were different across regions in both novice adolescents and novice adults (Supplementary Table 6).

Cortical activity during behavior reflects both age- and learning-induced effects.

Earlier version: A. Training and recording schedule for novice mice, compared to expert mice. Created with BioRender.com. B. 3D-Reconstruction of recording sites in novice adolescent (n = 6; grey) and novice adult (n = 6; black) mice. Bottom: spike-depth of excitatory tone-responsive L5/6 adolescent (n = 107; grey) and adult (n = 177; black) neurons. C. Normalized FR and lick rate (LR) PSTH from -200ms to 600ms after tone-onset in adolescents (grey) and adults (black). Average +-sem. D. Single neuron data from novice adolescent mice. Left: Heat map of the FR per trial from one example neuron sorted by trial outcomes. Center: the AUC of the neuron from the left for the easy and hard stimulus pairs (light and dark blue, respectively). Right: Average (+-SEM) AUC of all neurons in the novice group (n = 140 neurons). E-G. Same as ‘D’ for novice adult (n = 186 neurons), expert adolescents (n = 455 neurons; Easy vs hard), and expert adults (n = 604 neurons; Easy vs hard.). H. Linear regression analysis between the average AUC per recording and the behavioral d’ during the recording (the correlation and p values are indicated for each plot). I. Same as ‘I’ for adult mice.

Revised version: A. Training and recording schedule for novice mice, compared to expert mice. Created with BioRender.com. B. 3D-Reconstruction of recording sites in novice adolescent (n = 6; grey) and novice adult (n = 6; black) mice. Bottom: spike-depth of excitatory tone-responsive L5/6 adolescent (n = 130; grey) and adult (n = 186; black) neurons. C. Normalized FR and lick rate (LR) PSTH from −200ms to 600ms after tone-onset in adolescents (grey) and adults (black). Average +-sem. D. Single neuron data from novice adolescent mice. Left: Heat map of the FR per trial from one example neuron sorted by trial outcomes. Center: the AUC of the neuron from the left for the easy and hard stimulus pairs (light and dark blue, respectively). Right: Average (+-SEM) AUC of all neurons in the novice group (n = 140 neurons). E-G. Same as ‘D’ for novice adult (n = 186 neurons), expert adolescents (n = 455 neurons; Easy vs hard), and expert adults (n = 604 neurons; Easy vs hard.). H. Linear regression analysis between the average AUC per recording and the behavioral d’ during the recording (the correlation and p values are indicated for each plot). I. Same as ‘I’ for adult mice.

Earlier version: To reveal learning-related differences, we compared novice and expert mice of similar ages. We computed single-neuron AUCs for Go and No-Go trials (here we considered all trial outcomes to reliably assess differences between novice and expert mice, since expert mice mostly performed CR and novice mice mostly performed FA to No-Go tones), in both easy and hard stimulus pairs (Fig. 6d-g). The onset-latency for AUC discrimination was significantly different between adolescents and adults, but not between novice and expert mice, nor between the easy and the hard task. The maximal AUC discriminability was higher, and the duration longer in adults compared to adolescents, between experts and novice, and between the easy and hard tasks (Table 3). These results suggest that plasticity is reflected in the interaction of age and learning.

Revised version: To reveal learning-related differences, we compared neuronal responses in novice versus expert mice in both age-groups. We computed single-neuron AUCs for Go and No-Go trials (here we did not distinguish between stimulus and choice and computed AUCs for all trial outcomes, since expert mice mostly performed CR and novice mice mostly performed FA to No-Go tones), in both easy and hard task (Fig. 6d-g). The onset-latency for AUC discrimination was significantly different between adolescents and adults (shorter in adults), but not between novice and expert mice, nor between the easy and the hard task (Table 3). The maximal AUC discriminability was higher in adults compared to adolescents. It was also higher in experts than in novice, and in the easy compared to the hard task (Table 3). Similarly, the duration of discriminability was longer in adults and after learning (Table 3). There was a significant interaction between age and learning, as well learning and task difficulty in the maximal discriminability (Table 3). The onset-latency, duration of discriminability and maximal discriminability, were not significantly different among the AUDd, AUDp, AUDv and TEa in the four experimental groups (Fig. S6-2).

To compare neuronal discriminability and behavior more explicitly, we analyzed the correlations between neuronal decoding and behavioral performance using linear regression. In the easy task, both age groups showed high correlation between neuronal discriminability and behavioral discriminability (Fig. 6h,i; light blue). In the hard task, only adult neurons had significant correlations between neuronal discriminability and behavioral discriminability (Fig. 6h,i; dark blue).

Effects of age and learning on tuning properties in passively listening mice

Earlier version: Auditory discrimination learning has been shown to induce changes in the tuning properties of sounds in the ACx [18, 19, 21, 23–28]. To study whether learning-induced plasticity in ACx is distinct in adolescents and adults, we recorded from the same mice described above under passive listening condition, by simply extending the recording following the engaged session (Fig. 7a, Supplementary Table 5; Adolescents recordings = 4; Adults recordings = 4). These data were compared to a group of novice mice, recorded under passive listening following their engaged session (Supplementary Table 5; adolescent recordings =6; adult recordings =6).

Revised version: Auditory learning has been shown to induce changes in the tuning properties of sounds in the ACx [18, 19, 21, 23–28]. To study whether learning-induced plasticity in ACx is distinct in adolescents and adults, we recorded from the same mice described above under passive listening conditions by simply extending the recording following the engaged session (Fig. 7a, Supplementary Table 7; Adolescents recordings = 4; Adults recordings = 4). These data were compared to data from a group of novice mice, recorded under passive listening conditions following their own engaged session (Supplementary Table 7; novice adolescent recordings =6; novice adult recordings =6).

Added in the revised version: Adult mice show greater plasticity after learning.

Removed in the revised version: Unlike adults, neuronal tuning properties of adolescents do not change after learning

A. Schematic showing that for the passive listening protocol, we continued our recording following the session of the engaged task (i.e. in satiated mice) by removing the waterspout. Created with BioRender.com. B. Example raster plot of a neuron from an adolescent mouse (top) and an adult mouse (bottom). C. FRA’s of the neurons shown in ‘B’. D. Distribution of best frequencies in our dataset. Values are normalized firing rates calculated at 62 dB SPL. Matrices are sorted by BF for clarity. Dotted line marks the decision boundary. E. Tuning bandwidth at 62 dB SPL of neurons in adolescents and adults. Side by side comparisons of novice versus experts. (adolescents p = 0.0882, adults p = 0.0001, Kruskal Willis Test after Tukey-Kramer correction for multiple comparisons). F. Same as E. for Population sparseness (adolescents p = 0.9549, adults p = 0.0013, Kruskal Willis Test after Tukey-Kramer correction for multiple comparisons). G Same as E. for the distance (in octaves) between the best-frequency of each neuron to the easy Go-stimulus (adolescents p = 0.0816, adults p = 0.6391, Kruskal Willis Test after Tukey-Kramer correction for multiple comparisons). H Same as E. for the average neuronal d’ of frequencies in the learned frequency spectrum (adolescents p = 0.1627, adults p = 0.0026, Kruskal Willis Test after Tukey-Kramer correction for multiple comparisons).

Added in the revised version: We characterized the response profile of ACx neurons using a ‘frequency-response area’ protocol composed of twenty pure tones (4-40 kHz, spaced at 0.1661 octaves steps), each played for 100ms at five different attenuations (72-32 dB SPL). We limited our analysis to significant excitatory units only, determined as being auditory responsive in a window from tone onset to 50ms after tone offset at 62 dB SPL. We then focused our tuning analysis on responses to 62 dB SPL because this sound level was close to the intensity used during the behavioral task (72 dB SPL), while still avoiding potential ceiling effects observed at higher SPLs in recordings of mice not engaged in the task.

Earlier version: We characterized the response profile of ACx neurons using a ‘frequency-response area’ protocol composed of twenty pure tones (4-40 kHz, spaced at 0.1661 octaves steps), each played for 100ms at five different attenuations (72-42 dB SPL). We limited our analysis to significant excitatory units only, determined as being auditory responsive in a window from tone onset to 50ms after tone offset at 62 dB SPL. Our dataset included four groups. The two adolescent groups were: 1) adolescent novice, and 2) adolescent experts, with 92 and 80 neurons, respectively. The two adult groups were: 3) adult novices, and 4) adult experts, with 123 and 84 neurons, respectively. Neurons in all groups responded to pure tone as expected from previous studies in adult mice, including classical V-shaped frequency response areas [22, 29–32]. A representative example from a responsive neuron of the adolescent expert group and one from the adult expert group are shown in Figure 7b-c. We next compared the firing properties of different sub-regions (TEa and AUDd were excluded since less than 20 units responded to the FRA protocol after task engagement, see Supplementary Table 5). We found that the firing properties were distinct between adolescent and adult neurons, in both AUDp and AUDv. However there were no significant differences between AUDp and AUDv in either group (Supplementary Table 6). The peak of frequency tuning was heterogeneous across the frequency range with no clear trends of learning or age (Fig. 7d). To quantify these responses, we calculated tuning width (Fig. 7e), population sparseness (Fig. 7f), the distance between the best frequency of the neurons and the Go frequency of the easy task (Fig. 7g), and single neuron discriminability (Fig. 7h). As expected from previous work, learning significantly changed neuronal tuning properties in adults (Fig. 7e-h, ‘Adults’). Surprisingly, however, learning had no significant effects on any of the single neuron parameters that we tested in the adolescent group (Fig. 7e-h, ‘Adolescents’ In adults, these results remained consistent when AUDp and AUDv were analyzed separately (Fig. S7-1). In contrast, adolescent results were more heterogeneous: AUDp exhibited changes in tuning alongside higher discriminability, while AUDv showed no detectable differences. Thus, despite the notion that adolescents’ brains may be more plastic than adults, we found no clear evidence for learning-induced plasticity in basic tuning properties of the adolescent mice following our protocol (Table 5, Fig. S6-1). These results underscore that learning-induced plasticity manifests differently in adolescents and adults.

Revised version: Our dataset included four groups. The two adolescent groups were: 1) adolescent novices, and 2) adolescent experts, with 92 and 80 neurons, respectively. The two adult groups were: 3) adult novices, and 4) adult experts, with 123 and 84 neurons, respectively. Neurons in all groups responded to pure tone as expected from previous studies in adult mice, including classical V-shaped frequency response areas [22, 29–32]. A representative example from a responsive neuron of an adolescent expert and one from an adult expert are shown in Figure 7b,c. To assess differences in pure-tone responses between auditory subregions, we quantified the firing properties of AUDp versus AUDv (Supplementary Table 8; TEa and AUDd were excluded since fewer than 20 units responded to the FRA protocol after task engagement, see Supplementary Table 7). We found that the latency to peak was slower in adolescent AUDp. The FWHM and minimal latency were slower in the adolescents in both AUDp and AUDv. In the adolescent AUDv, additional parameters (like spontaneous- & evoked-FR, trial-responsiveness, life time sparseness), were different than those in adults. There were no significant differences between AUDp and AUDv within age-groups (Supplementary Table 8).

Added in the revised version: Next, we tested the learning-related effects among the different groups. The peak of frequency tuning was heterogeneous across the frequency range with no clear trends of learning or age (Fig. 7d). To quantify these responses, we calculated tuning width (Fig. 7e), population sparseness (Fig. 7f), the distance between the best frequency of the neurons and the Go frequency of the easy task (Fig. 7g), and single neuron discriminability (Fig. 7h). As expected from previous work, learning significantly changed (3 out of 4) neuronal tuning properties in adults (Fig. 7e-h, ‘Adults’). Surprisingly, however, learning had no significant effects on any of these single neuron parameters in the adolescent group (Fig. 7e-h). To test this at specific brain regions we analyzed AUDp and AUDv separately, which revealed specificity in learning-induced changes. Learning significantly affected 3 out of the 4 parameters in AUDp of adolescents (Fig. S7-1a-d). However, neurons from AUDv of adolescents showed no learning-induced changes (Fig. S7-1e-h). Notably, adult learning-induced plasticity was stronger than that in adolescence in all parameters besides the ‘BF to Go’ (Fig. S7-1). Thus, despite the notion that adolescents’ brains may be more plastic than adults, we found generally stronger learning-induced plasticity in basic tuning properties in adults. These results underscore that learning-induced plasticity manifests differently in adolescents and adults.

Discussion

The term “stormy adolescence” aptly captures what we know from psychology, neuroscience, and biology—that adolescence is a period of turbulent development, which shapes the individual’s future behavior and health[33, 34]. While numerous behavioral phenomena of adolescence have been studied, particularly in humans, the neural underpinnings of these remain largely unknown[35]. In this study, we focused on (auditory) cortical neurons of the adolescent brain of mice while they were engaged in an auditory discrimination task. We found numerous differences in behavior, single neuron activity, and population encoding between adolescents and adults.

Adolescence behavior – perceptual and cognitive differences

Earlier version: We utilized the Educage platform to compare between adolescents and adults[18]. The Educage provides an optimal environment for training young mice; its automated nature minimizes human handling, which increases efficiency and reduces to a minimum any interference during the training process. Using the Educage, we found that procedural learning was intact in pre-adolescent mice (P20-P25), and behavioral differences were found only after the mice learned the task during the discrimination phase (P30-37). Compared to adults, adolescent mice scored lower on the hard versions of the task (Fig. 1i-j). Interestingly, our data suggests that the lower performance in the hard task in adolescents was not due to a deficiency in perceptual sensitivity (Fig. 2b), but rather to a cognitive deficit – namely, response bias (Fig. 2c). This result is perhaps not surprising since higher response bias is a cognitive phenomenon that has been well documented in adolescence as a critical period of inhibitory control[36]. Indeed, additional cognitive control mechanisms are still developing during this age [37–39].

Revised version: We utilized the Educage platform to compare between adolescents and adults[18]. The Educage provides an optimal environment for training young mice; its automated nature minimizes human handling, which increases efficiency and reduces to a minimum any interference during the training process. Using the Educage, we found that procedural learning was intact in pre-adolescent mice (P23-P30), and behavioral differences were found only after the mice learned the task during the hard discrimination phase (P30-37). Compared to adults, adolescent mice scored lower on the hard versions of the task (Fig. 1i-j). Interestingly, our data suggests that the lower performance in the hard task in adolescents was not due to a deficiency in perceptual sensitivity (Fig. 2c), but rather to a cognitive deficit – namely, lick bias (Fig. 2d). This result is perhaps not surprising since higher lick bias is a cognitive phenomenon that has been well documented in adolescence as a maturation period of inhibitory control[36]. Indeed, additional cognitive control mechanisms are still developing during this age [37–39].

Earlier version: Another manifestation of weaker cognitive control mechanisms is the noisy behavioral performance of adolescents (Fig. 2h-j). While this noise may reflect higher plasticity and potentially serve cognitive flexibility[6, 40], it poses a disadvantage for the precision and consistency required in our task. This noise may also arise from changes in pubertal hormones, which we did not manipulate[41]. Notably, few studies have explicitly characterized auditory discrimination of adolescents in rodents; the majority of these have been carried out in gerbils using an amplitude modulation (AM) detection task [4, 26, 42]. Despite the different task, many of our age-related findings are aligned with the findings in gerbils, including lower performance in adolescents, and greater variance in adolescent behavior[4]. However, some findings differ. For example, the response bias which was different in our data, was similar between young and adult gerbils in an AM modulation detection task. This difference may arise from differences in task design, as well as the severity of the punishment used in the task (here, we used a milder punishment). Together, our data support a view that cognitive factors, and to a lesser extent perception of pure tones per se, are the primary contributors to the poorer performance observed in adolescent mice when performing a sound discrimination task.

Revised version: Another manifestation of weaker cognitive control mechanisms is the noisy behavioral performance of adolescents (Fig. 2h-j). While this noise may reflect higher plasticity and potentially serve cognitive flexibility[6, 40], it poses a disadvantage for the precision and consistency required in our task. This noise may also arise from changes in pubertal hormones, which we did not manipulate[41]. Notably, few studies have explicitly characterized auditory discrimination of adolescents in rodents; the majority of these have been carried out in gerbils using an amplitude modulation (AM) detection task [4, 26, 42]. Despite the different task, many of our age-related findings are aligned with the findings in gerbils, including lower performance in adolescents, and greater variance in adolescent behavior[4]. However, some findings differ. For example, the lick bias which was different in our data, was similar between young and adult gerbils in an AM modulation detection task. This difference may arise from differences in task design, as well as the severity of the punishment used in the task (here, we used a milder punishment). Here, our data support a view that cognitive factors, and to a lesser extent perception of pure tones per se, are the primary contributors to the poorer performance observed in adolescent mice when performing a pure-tone sound discrimination task.

Representation of stimulus and choice in ACx still undergoes maturation in adolescence

Earlier version: Motivated by the abovementioned behavioral differences, we tested how single neurons in ACx—whose involvement in task performance we confirmed via optogenetic silencing in adult mice (Fig. S3-2)—represented sounds (by comparing Hits vs FA) and choices (by comparing FA to CR). While some studies suggest that ACx is not required for simple auditory-guided behaviors [43], others highlight its role in pure tone perception and task execution [19, 20, 44, 45]. While cortical representations of sounds in ACx could be the basis of perceptual constraints, distinct processing of choices would reflect that more cognitive mechanisms are involved. We found that cortical representations of sounds and choices in mice engaged in the task were still immature as compared to adults, e.g., responses in adolescents were more sluggish and less informative (Figs. 4-6). These differences can be a result of age-related differences, and/or the result of a combined ‘age x learning’ effect.

Revised version: Motivated by the abovementioned behavioral differences, we tested how single neurons in ACx represented sounds (by comparing Hits vs FA) and choices (by comparing FA to CR). While cortical representations of sounds in ACx could be the basis of perceptual constraints, distinct processing of choices would reflect that more cognitive mechanisms are involved. We found that cortical representations of sounds and choices in mice engaged in the task were still immature as compared to adults, e.g., responses in adolescents were more delayed and less informative (Figs. 4-6). These differences can be a result of age-related differences, and/or the result of a combined ‘age x learning’ effect.

Earlier version: We found modest age-related differences in spiking activity of neurons from layer 5/6 between adolescents and adults (Fig. 6d-g; Fig. 7). This limited difference is perhaps not surprising because the CP for tonal organization in ACx is already complete by the age at which we tested the animals (i.e., on P38 [9, 13]). In contrast, more pronounced differences emerged from the interaction of ‘age x learning’, specifically in mice that learned the task and were actively engaged in the behavior (Fig. 4e, f; Fig. 6d-g). This interaction is evident in the reduced single-neuron discriminability (AUC) during the hard task for expert adolescent mice (Fig. 6d vs. 6f). While this may seem counterintuitive, it is consistent with the idea that learning can drive complex reorganization in sensory representations. For example, we recently showed that learning leads to distinct and sometimes opposing effects on single-neuron responses with respect to whether the task is perceptually easy or difficult [46].

Revised version: We found only modest age-related differences in spiking activity of layer 5/6 neurons between adolescents and adults (Fig. 7). This limited difference is perhaps not surprising because the CP for tonal organization in ACx is already complete by the age at which we tested the animals (i.e., on P37-P42 [9, 13]). In contrast, more pronounced differences emerged from the interaction of ‘age x learning’, specifically in mice that learned the task and were actively engaged in the behavior (Fig. 4f; Fig. 6d-g). This interaction is evident in the lower single-neuron discriminability (AUC) during the hard task in expert adolescent mice (Fig. 6d vs. 6f).

Added in the revised version: While this may seem counterintuitive, it is compatible with the idea that learning can drive complex reorganization in sensory representations. For example, we recently showed that learning leads to distinct and sometimes opposing effects on single-neuron responses with respect to whether the task is perceptually easy or difficult [46]. It is also important to note that our recordings were restricted to infragranular layers (5/6) and did not sample supragranular (layers 2/3) neurons, which may exhibit distinct learning related plasticity [21]. As such, the developmental differences we report may reflect layer-specific properties of the deep output layers.

Earlier version: It is important to note that our recordings were restricted to infragranular layers (5/6) and did not sample supragranular layers (e.g., layers 2/3), which may exhibit distinct plasticity profiles [21]. As such, the developmental differences we report may reflect layer-specific properties of deep output layers, particularly their involvement in learning-related plasticity and communication with subcortical or frontal targets. We speculate that immature feedback (FB) to ACx may underlie these differences for several reasons. First, FB from higher cortical regions to more primary cortices has been proposed as a neural substrate for difficult perceptual discriminations[47, 48]. Second, FB activity (often referred to as ‘top down’) is thought to be the key neural pathway involved in cognitive control mechanisms[49]. FB is typically expressed in the late (>100ms) phase of single neurons activity, which was significantly different between adults and adolescents in our data (Fig. 4d). Furthermore, FB may act to stabilize or suppress sensory responses in ACx, reducing apparent single-neuron discriminability of hard stimuli after learning (Fig. 6f), while preserving or enhancing in the population discriminability such as we observed in Fig. 5.

Revised version: We speculate that immature feedback to ACx may underlie these differences for several reasons. First, feedback from higher cortical regions to more primary cortices has been proposed as a neural substrate for difficult perceptual discriminations[47, 48]. Second, feedback activity (often referred to as ‘top down’) is thought to be the key neural pathway involved in cognitive control mechanisms[49]. Feedback is typically expressed in the late (>100ms) phase of single neurons activity, which was significantly different between adults and adolescents across our measured parameters (onset-latency, duration of discriminability and maximal discriminability; Fig. 4g, Table 2, Table 3). The precise source of feedback that contributes to perceptual or cognitive differences remains to be determined, but likely candidates include regions that project to ACx [50, 51], such as medial prefrontal cortex[52, 53], orbitofrontal cortex[54], or anterior cingulate[55]. Testing the causal role of these regions during adolescence will be an important direction for future research.

Removed in the revised version: The precise source of FB that contributes to perceptual or cognitive differences remains to be determined, but likely candidates include regions that project to ACx [50, 51], such as medial prefrontal cortex[52, 53], orbitofrontal cortex[54], or anterior cingulate[55]. Testing the causal role of these regions during adolescence will be an important direction for future research.

Earlier version: Cortical coding is performed by populations of neurons, which can efficiently encode stimuli on a trial-by-trial basis[56–60]. Population coding, as measured here by decoding accuracy, largely recapitulated the results of the single neuron data. Specifically, cortical populations in adults encoded the different tones better and faster; particularly in the hard task (Fig. 5a-c). Interestingly, by testing the quality (i.e., separation) of the decoding performance, we revealed a correlation between decoding in the easy and hard tasks in adult (Fig. 5d). This result suggests that the cortical network uses shared representation for decoding the easy and hard tasks, possibly by the same neurons. This is consistent with our recent work showing that indeed the same neurons in ACx respond to the same stimulus differently, depending on if the mouse is engaged in discriminating an easy sound pair or a hard sound pair[46]. The shallower slopes in adolescence suggest a less efficient decoding scheme that may arise from different populations of neurons encoding the stimuli in the easy and hard versions of the task.

Revised version: Cortical coding is performed by populations of neurons, which can efficiently encode stimuli on a trial-by-trial basis[56–60]. Population coding, as measured here by decoding accuracy, largely recapitulated the results of the single neuron data. Specifically, cortical populations in adults encoded the different tones better and faster; particularly in the hard task (Fig. 5a-c). Interestingly, by testing the quality (i.e., separation) of the decoding performance, we revealed a correlation between decoding in the easy and hard tasks in adult (Fig. 5d). While this result reflects the collapsed information of stimulus and choice, it suggests that the cortical network uses shared representation for decoding the easy and hard tasks, possibly by the same neurons. This is consistent with our recent work showing that indeed the same neurons in ACx respond to the same stimulus differently, depending on if the mouse is engaged in discriminating an easy sound pair or a hard sound pair[46]. The shallower slopes in adolescence suggest a less efficient decoding scheme that may arise from different populations of neurons encoding task-related information.

Earlier version: Our optogenetic silencing of ACx in expert adult mice resulted in a marked suppression of licking behavior (Fig. S3-2), consistent with the idea that ACx contributes to task performance. However, the strength of the effect raises the possibility that the manipulation may have impacted broader aspects of behavior—such as arousal, motivation, or motor readiness—rather than selectively disrupting auditory discrimination. Recent work [61] suggests that ACx plays a critical role during learning by supporting higher-order computations but may become less essential for task execution once mice reach expert performance. These findings, together with previous reports [19, 20], indicate that the necessity of ACx for auditory-guided behavior may depend on task structure, training stage, and the exact nature of the manipulation. The robust decrease in lick rate and performance observed in our experiment may reflect the intensity, specificity, and timing of the silencing. The use of GtACR2, a highly potent silencer, may have further amplified the effect. We cannot rule out the possibility that silencing the output from ACx also disrupted the flow of information to downstream areas, thereby exerting a stronger influence on perception and behavior. Although the injection was targeted to the ACx, off-target effects are also likely to have contributed. Future experiments using graded or temporally specific perturbations will be useful to more precisely defining the contribution of ACx across learning and development.

Revised version: Our optogenetic silencing of ACx in expert adult mice resulted in a marked suppression of licking behavior (Fig. S3-2), consistent with the idea that ACx contributes to task performance. However, the strength of the observed effect raises the possibility that the manipulation may have impacted broader aspects of behavior—such as arousal, motivation, or motor readiness—rather than selectively disrupting auditory discrimination. Recent work [61] suggests that ACx plays a critical role during learning by supporting higher-order computations but may become less essential for task execution once mice reach expert performance. These findings, together with previous reports [19, 20], indicate that the necessity of ACx for auditory-guided behavior may depend on task structure, training stage, and the exact nature of the manipulation. The robust decrease in lick rate and performance observed in our experiment may reflect the intensity, specificity, and timing of the silencing. The use of GtACR2, a highly potent silencer, may have further amplified the effect. We cannot rule out the possibility that silencing the output from ACx also disrupted the flow of information to downstream areas, thereby exerting a stronger influence on perception and behavior. Although the injection was targeted to the ACx, off-target effects are also likely to have contributed. Future experiments using graded or temporally specific perturbations will be useful to more define more precisely the contribution of ACx across learning and development.

Learning in the adolescent brain

Earlier version: Cortical representations of pure tones have a critical period (CP) between P12-P15[13, 16, 17]. Indeed, we chose to study animals at P38-P40 on a pure tone discrimination task precisely to avoid the hypersensitivity associated with this CP. But this may be unavoidable, as recent work has shown that ACx also has a later CP for more complex sounds features[16, 17]. Those late CP’s overlap with our training schedule and the underlying biological manifestation of this CP is expected to affect learning related plasticity in general. The underlying mechanism of the late CP to complex sounds has been associated with dynamic changes of inhibitory circuits[62, 63]. Thus, immature inhibitory circuits remain sensitive in adolescents, and could affect how pure tones are represented, despite the closure of the classic tonal CP. Notably, sound discrimination learning changes tonal representations even in adult mice, well after the closure of CP’s for both simple and complex sounds [18, 24, 46]. We hypothesize, therefore, that the differences in performance on more difficult discriminations between adolescents and adults are shaped primarily by age-specific learning-related plasticity, rather than by differences in pure tone representation per se.

Revised version: Cortical representations of pure tones have a critical period (CP) between P12-P15[13, 16, 17]. Indeed, we chose to study animals at P37-P42 on a pure tone discrimination task precisely to avoid the hypersensitivity associated with this CP. But this may be unavoidable, as recent work has shown that ACx also has a later CP for more complex sounds features[16, 17]. Those late CP’s overlap with our training schedule and the underlying biological manifestation of this CP is expected to affect learning related plasticity in general. The underlying mechanism of the late CP to complex sounds has been associated with dynamic changes of inhibitory circuits[62, 63]. Thus, immature inhibitory circuits remain sensitive in adolescents, and could affect how pure tones are represented, despite the closure of the classic tonal CP. Notably, sound discrimination learning changes tonal representations even in adult mice, well after the closure of CP’s for both simple and complex sounds [18, 24, 46]. We hypothesize, therefore, that the differences in performance on more difficult discriminations between adolescents and adults are shaped primarily by age-specific learning-related plasticity, rather than by differences in pure tone representation per se.

Earlier version: Taken together, our findings show that cortical responses in adolescents differ from those in adults, despite both groups learning the same pure tone discrimination task under identical training protocols. Although the CPs for pure tones are considered closed by this age, learning to discriminate between tones still varies across developmental stages. It remains unclear how later CPs—such as those shaping responses to amplitude modulation (AM) or frequency-modulated (FM) sweeps through inhibitory circuits—contribute to these differences. Likewise, the role of ongoing maturation in cognitive control, including the development of feedback pathways, warrants further investigation. Nevertheless, it is evident that developmental processes in the adolescent brain continue to influence how even simple sounds like pure tones are encoded. The adolescent-specific plasticity observed here directly reflects how perceptual and cognitive features are integrated within the auditory cortex (ACx).

Revised version: Taken together, our findings show that cortical responses in adolescents differ from those in adults, despite both groups learning the same pure tone discrimination task under identical training protocols. Although the CPs for pure tones are considered closed by this age, learning to discriminate between tones still varies across developmental stages. It remains unclear how later CPs, such as those shaping responses to amplitude modulation or frequency-modulated sweeps through inhibitory circuits, contribute to these differences. Likewise, the role of ongoing maturation in cognitive control, including the development of feedback pathways, warrants further investigation. Nevertheless, it is evident that developmental processes in the adolescent brain continue to influence how even simple sounds are encoded through learning. The adolescent-specific plasticity observed here directly reflects on how perceptual and cognitive features are integrated within the auditory cortex.